Linking the evolution of two prefrontal brain regions to social and foraging challenges in primates

Author Response

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public Review):

The present study provides a phylogenetic analysis of the size prefrontal areas in primates, aiming to investigate whether relative size of the rostral prefrontal cortex (frontal pole) and dorsolateral prefrontal cortex volume vary according to known ecological or social variables.

I am very much in favor of the general approach taken in this study. Neuroimaging now allows us to obtain more detailed anatomical data in a much larger range of species than ever before and this study shows the questions that can be asked using these types of data. In general, the study is conducted with care, focusing on anatomical precision in definition of the cortical areas and using appropriate statistical techniques, such as PGLS. That said, there are some points where I feel the authors could have taken their care a bit further and, as a result, inform the community even more about what is in their data.

We thank the reviewer for this globally positive evaluation of our work, and we appreciate the advices to improve our manuscript.

The introduction sets up the contrast of 'ecological' (mostly foraging) and social variables of a primate's life that can be reflected in the relative size of brain regions. This debate is for a large part a relic of the literature and the authors themselves state in a number of places that perhaps the contrast is a bit artificial. I feel that they could go further in this. Social behavior could easily be a solution to foraging problems, making them variables that are not in competition, but simply different levels of explanation. This point has been made in some of the recent work by Robin Dunbar and Susanne Shultz.

Thank you for this constructive comment, and we acknowledge that the contrast between social vs ecological brain is relatively marginal here. Based also on the first remark by reviewer 3, we have reformulated the introduction to emphasize what we think is actually more critical: the link between cognitive functions as defined in laboratory conditions and socio-ecological variables measured in natural conditions. And the fact that here, we use brain measures as a potential tool to relate these laboratory vs natural variables through a common scenario. Also, we were already mentioning the potential interaction between social and foraging processes in the discussion, but we are happy to add a reference to recent studies by S. Shultz and R. Dunbar (2022), which is indeed directly relevant. We thank the reviewer for pointing out this literature.

In a similar vein, the hypotheses of relating frontal pole to 'meta-cognition' and dorsolateral PFC to 'working memory' is a dramatic oversimplification of the complexity of cognitive function and does a disservice to the careful approach of the rest of the manuscript.

We agree that the formulation of which functions we were attributing to the distinct brain regions might not have been clear enough, but the functional relation between frontal pole and metacognition in the one hand, and DLPFC and working memory on the other hand, have been firmly established in the literature, both through laboratory studies and through clinical data. Clearly, no single brain region is necessary and sufficient for any cognitive operation, but decades of neuropsychology have demonstrated the differential implication of distinct brain regions in distinct functions, which is all we mean here. We have made a specific point on that topic in the discussion (cf p. 16). We have also reformulated the introduction to clarify that, even if the relation between these regions and their functions (FP/ metacognition; DLPFC/ working memory) was clear in laboratory conditions, it was not clear whether this mapping could be used for real life conditions. And therefore whether that simplification was somehow justified beyond the lab (and the clinics), and whether these neuro-cognitive concepts could be applied to natural conditions, are indeed critical questions that we wanted to address. The central goal of the present study was precisely to evaluate the extent to which this brain/cognition relation could be used to understand more natural behaviors and functions, and we hope that it appears more clearly now.

One can also question the predicted relationship between frontal pole meta-cognition and social abilities versus foraging, as Passingham and Wise show in their 2012 book that it is frontal pole size that correlates with learning ability-an argument that they used to relate this part of the brain to foraging abilities. I would strongly suggest the authors refrain from using such descriptive terms. Why not simply use the names of the variables actually showing significant correlations with relative size of the areas?

We basically agree with the reviewer, and we acknowledge the lack of clarity in the introduction of the previous manuscript. There were indeed lots of ambiguity in what we were referring to as ‘function’, associated with a given brain region. « Function » referred to way to many things! We have reformulated the introduction not only to clarify the different types of functions that were attributed to distinct brain regions in the literature but also to clarify how this study was addressing the question: by trying to articulate concepts from neuroscience laboratory studies with concepts from behavioral ecology and evolution using intuitive scenarios. We hope that the present version of the introduction makes that point clearer.

The major methodological judgements in this paper are of course in the delineation of the frontal pole and dorsolateral prefrontal cortex. As I said above, I appreciate how carefully the authors describe their anatomical procedure, allowing researchers to replicate and extend their work. They are also careful not to relate their regions of interest to precise cytoarchitectonic areas, as such a claim would be impossible to make without more evidence. That said, there is a judgement call made in using the principal sulcus as a boundary defining landmark for FP in monkeys and the superior frontal sulcus in apes. I do not believe that these sulci are homologous. Indeed, the authors themselves go on to argue that dorsolateral prefrontal cortex, where studied using cytoarchitecture, stretches to the fundus of principal sulcus in monkeys, but all the way to the inferior frontal sulcus in apes. That means that using the fundus of PS is not a good landmark.

We thank the reviewer for his kind remarks on our careful descriptions. But then, it is not clear whether our choice of using the principal sulcus as a boundary for FP in monkeys vs the superior frontal sulcus in apes is actually a judgement call. First, and foremost, there is no clear and unambiguous definition of what should be the boundaries of the FP. By contrast with cytoarchitectonic maps, but clearly this is out of reach here. In humans and great apes we used Bludau et al 2014 (i.e. sup frontal sulcus), and in monkeys, we chose a conservative landmark that eliminated area 9, which is traditionally associated with the DLPFC (Petrides, 2005; Petrides et al, 2012; Semendeferi et al, 2001).

Of course, any definition will attract criticism, so the best solution might be to run the analysis multiple times, using different definitions for the areas, and see how this affects results.

Indeed, functional maps indicate that dorsal part of anterior PFC in monkeys is functionally part of FP. But again, cytoarchitectonic maps also indicate that this part of the brain includes BA 9, which is traditionally associated with DLPFC (Petrides, 2005; Petrides et al, 2012). As already pointed out in the discussion, there is a functional continuum between FP and DLPFC and our goal when using PS as dorsal border was to be very conservative and to exclude the ambiguous area. But we agree with the reviewer that given that this decision is arbitrary, it was worth exploring other definitions of the FP volume. So, we did complete a new analysis with a less conservative definition of the FP, to include this ambiguous dorsal area, and it is now included in the supplementary material. Maybe as expected, including the ambiguous area in the FP volume shifted the relation with socio-ecological variables towards the pattern displayed by the DLPFC (ie the influence of population density decreased). The most parsimonious interpretation of this results is that when extending the border of the FP region to cover a part of the brain which might belong to the DLPFC, or which might be somehow functionally intermediate between the 2, the specific relation of the FP with socio-ecological variables decreases. Thus, even if we agree that it was important to conduct this analysis, we believe that it only confirms the difficulty to identify a clear boundary between FP and DLPFC. Again, we have clearly explained throughout the manuscript that we admit the lack of precision in our definitions of the functional brain regions. In that frame, the conservative option seems more appropriate and for the sake of clarity, the results of the additional analysis of a FP volume that includes the ambiguous area is only included in the supplementary material.

If I understand correctly, the PGLS was run separately for the three brain measure (whole brain, FP, DLPFC). However, given that the measures are so highly correlated, is there an argument for an analysis that allows testing on residuals. In other words, to test effects of relative size of FP and DLPFC over and above brain size?

Generally, using residuals as “data” (or pseudo-data) is not recommended in statistical analyses. Two widely cited references from the ecological literature are:

Garcia-Berthou E. (2001) On the Misuse of Residuals in Ecology: Testing Regression Residuals vs. the Analysis of Covariance. Journal of Animal Ecology, 70 (4): 708-711.

Freckleton RP. (2002). On the misuse of residuals in ecology: regression of residuals vs. multiple regression. Journal of Animal Ecology 71: 542–545. https://doi.org/10.1046/ j.1365-2656.2002.00618.x.

The main reason for this recommendation is that residuals are dependent on the fitted model, and thus on the particular sample under consideration and the eventual significant effects that can be inferred.

In the discussion and introduction, the authors discuss how size of the area is a proxy for number of neurons. However, as shown by Herculano-Houzel, this assumption does not hold across species. Across monkeys and apes, for instance, there is a different in how many neurons can be packed per volume of brain. There is even earlier work from Semendeferi showing how frontal pole especially shows distinct neuron-to-volume ratios.

We appreciate the reviewer’s comment, but the references to Herculano-Houzel that we have in mind do indicate that the assumption is legitimate within primates.

Herculano-Houzel et al (2007) show that the neuronal density of the cortex is well conserved across primate species (but only monkeys were studied). The conclusion of that study is that using volumes as a proxy for number of neurons, as a measure of computational capacity, should be avoided between rodents and primates (and as they showed later, even more so with birds, for which neuronal density is higher). BUT within primates, since neuronal densities are conserved, volume is a good predictor of number of neurons. Gabi et al (2016) provide evidence that the neuronal density of the PFC is well conserved between humans and non-human primates, which implies that including humans and great apes in the comparison is legitimate. In addition, the brain regions included in the analysis presumably include very similar architectonic regions (e.g. BA 10 for FP, BA 9/46 for DLPFC), which also suggests that the neuronal density should be relatively well conserved across species. Altogether, we believe that there is sufficient evidence to support the idea that the volume of a PFC region in primates is a good proxy for the number of neurons in that region, and therefore of its computational capacity.

Semendeferi and colleagues (2001) pointed out some differences in cytoarchitectonic properties across parts of the FP and discussed how these properties could 1) be used to identify area 10 across species 2) be associated with distinct computational properties, with the idea that thicker ‘cell body free’ layers would leave more space for establishing connections (across dendrites and axons). This pioneering work, together with more recent imaging studies on functional connectivity (e.g. Sallet et al, 2013) emphasize the critical contribution of connectivity pattern as a tool for comparative anatomy. But unfortunately, as pointed out in the discussion already, this is currently out of reach for us.

We acknowledge the limitations, and to be fair, the notion of computational capacity itself is hard to define operationally. Based on the work of Herculano-Houzel et al, average density is conserved enough across primates (including humans) to justify our approximation. We have tried to define our regions of interest using both anatomical and functional maps and, thanks to the reviewer’s suggestions, we even tried several ways to segment these regions. Functional maps in macaques and humans do not exactly match cytoarchitectonic maps, presumably because functions rely not only upon the cytoarchitectonics but also on connectivity patterns (e.g. Sallet et al, 2013).

In sum, we appreciate the reviewer’s point but feel that, given the current understanding of brain functions and the relative conservation of neuronal density across primate PFC regions, the volume of a PFC region seems to be reasonable proxy for its number of neurons, and therefore its computational capacity. We have added these points to the discussions, and we hope that the reader will be able to get a fair sense of how legitimate is that position, given the literature.

Overall, I think this is a very valuable approach and the study demonstrates what can now be achieved in evolutionary neuroscience. I do believe that they authors can be even more thorough and precise in their measurements and claims.

Reviewer #2 (Public Review):

In the manuscript entitled "Linking the evolution of two prefrontal brain regions to social and foraging challenges in primates" the authors measure the volume of the frontal pole (FP, related to metacognition) and the dorsolateral prefrontal cortex (DLPFC, related to working memory) in 16 primate species to evaluate the influence of socio-ecological factors on the size of these cortical regions. The authors select 11 socio-ecological variables and use a phylogenetic generalized least squares (PGLS) approach to evaluate the joint influence of these socio-ecological variables on the neuro-anatomical variability of FP and DLPFC across the 16 selected primate species; in this way, the authors take into account the phylogenetic relations across primate species in their attempt to discover the influence of socio-ecological variables on FP and DLPF evolution.

The authors run their studies on brains collected from 1920 to 1970 and preserved in formalin solution. Also, they obtained data from the Mussée National d´Histoire Naturelle in Paris and from the Allen Brain Institute in California. The main findings consist in showing that the volume of the FP, the DLPFC, and the Rest of the Brain (ROB) across the 16 selected primate species is related to three socio-ecological variables: body mass, daily traveled distance, and population density. The authors conclude that metacognition and working memory are critical for foraging in primates and that FP volume is more sensitive to social constraints than DLPFC volume.

The topic addressed in the present manuscript is relevant for understanding human brain evolution from the point of view of primate research, which, unfortunately, is a shrinking field in neuroscience.

We must not have been clear enough in our manuscript, because our goal is precisely not to separate humans from other primates. This is why, in contrast to other studies, we have included human and non-human primates in the same models. If our goal had been to study human evolution, we would have included fossil data (endocasts) from the human lineage.

But the experimental design has two major weak points: the absence of lissencephalic primates among the selected species and the delimitation of FP and DLPFC. Also, a general theoretical and experimental frame linking evolution (phylogeny) and development (ontogeny) is lacking.

We admit that lissencephalic species could not be included in this study because we use sulci as key landmarks. We believe that including lissencephalic primates would have introduced a bias and noise in our comparisons, as the delimitations and landmarks would have been different for gyrencephalic and lissencephalic primates. Concerning development, it is simply beyond the scope of our study.

Major comments.

1. Is the brain modular? Is there modularity in brain evolution?: The entire manuscript is organized around the idea that the brain is a mosaic of units that have separate evolutionary trajectories:

"In terms of evolution, the functional heterogeneity of distinct brain regions is captured by the notion of 'mosaic brain', where distinct brain regions could show a specific relation with various socio-ecological challenges, and therefore have relatively separate evolutionary trajectories".

This hypothesis is problematic for several reasons. One of them is that each evolutionary module of the brain mosaic should originate in embryological development from a defined progenitor (or progenitors) domain [see García-Calero and Puelles (2020)]. Also, each evolutionary module should comprise connections with other modules; in the present case, FP and DLPFC have not evolved alone but in concert with, at least, their corresponding thalamic nuclei and striatal sector. Did those nuclei and sectors also expand across the selected primate species? Can the authors relate FP and DLPFC expansion to a shared progenitor domain across the analyzed species? This would be key to proposing homology hypotheses for FP and DLPFC across the selected species. The authors use all the time the comparative approach but never explicitly their criteria for defining homology of the cerebral cortex sectors analyzed.

We do not understand what the referee is referring to with the word ‘module’, and why it relates to development. Same thing for the anatomical relation with subcortical structures. Yes, the identity of distinct functional cortical regions relies upon subcortical inputs during development, but clearly this is neither technically feasible, nor relevant here anyways.

We acknowledge, however, that our definition of functional regions was not precise enough, and we have updated the introduction to clarify that point. In short, we clearly do not want to make a strong case for the functional borders that we chose for the regions of interest here (FP and DLPFC), but rather use those regions as proxies for their corresponding functions as defined in laboratory conditions for a couple of species (rhesus macaques and humans, essentially).

Contemporary developmental biology has showed that the selection of morphological brain features happens within severe developmental constrains. Thus, the authors need a hypothesis linking the evolutionary expansion of FP and DLPFC during development. Otherwise, the claims form the mosaic brain and modularity lack fundamental support.

Once again, we do not think that our definition of modules matches what the reviewer has in mind, i.e. modules defined by populations of neurons that developed together (e.g. visual thalamic neurons innervating visual cortices, themselves innervating visual thalamic neurons). Rather, the notion of mosaic brain refers to the fact that different parts of the brain are susceptible to distinct (but not necessarily exclusive) sources of selective pressures. The extent to which these ‘developmental’ modules are related to ‘evolutionary’ modules is clearly beyond the scope of this paper.

Our goal here was to evaluate the extent to which modules that were defined based on cognitive operations identified in laboratory conditions could be related (across species) to socio-ecological factors as measured in wild animals. Again, we agree that the way these modules/ functional maps were defined in the paper were confusing, and we hope that the new version of the manuscript makes this point clearer.

Also, the authors refer most of the time to brain regions, which is confusing because they are analyzing cerebral cortex regions.

We do not understand why the term ‘brain’ is more confusing than ‘cerebral cortex’, especially for a wide audience.

1. Definition and delimitation of FP and DLPFC: The precedent questions are also related to the definition and parcellation of FP and DLPFC. How homologous cortical sectors are defined across primate species? And then, how are those sectors parcellated?

The authors delimited the FP:

"...according to different criteria: it should match the functional anatomy for known species (macaques and humans, essentially) and be reliable enough to be applied to other species using macroscopic neuroanatomical landmarks".

There is an implicit homology criterion here: two cortical regions in two primate species are homologs if these regions have similar functional anatomy based on cortico-cortical connections. Also, macroscopic neuroanatomical landmarks serve to limit the homologs across species.

This is highly problematic. First, because similar function means analogy and not necessarily homology [for further explanation see Puelles et al. (2019); García-Cabezas et al. (2022)].

We are not sure to follow the Reviewer’s point here. First, it is not clear what would be the evolutionary scenario implied by this comment (evolutionary divergence followed by reversion leading to convergence?). Second, based on the literature, both the DLPFC and the FP display strong similarities between macaques and humans, in terms of connectivity patterns (Sallet et al, 2013), in terms of lesion-induced deficit and in terms of task-related activity (Mansouri et al, 2017). These criteria are usually sufficient to call 2 regions functionally equivalent. We do not see how this explanation is "highly problematic" as it is clearly the most parsimonious based on our current knowledge.

Second, because there are several lissencephalic primate species; in these primates, like marmosets and squirrel monkeys, the whole approach of the authors could not have been implemented. Should we suppose that lissencephalic primates lack FP or DLPFC?

We understand neither the reviewer’s logic, nor the tone. We understand that the reviewer is concerned by the debate on whether some laboratory species are more relevant than others for studying the human prefrontal cortex, but this is clearly not the objective of our work. As explained in the manuscript, we identified FP and DLPFC based on functional maps in humans and laboratory monkeys (macaques), and we used specific gyri as landmarks that could be reliably used in other species. And, as rightfully pointed out by reviewer 1, this is in and off itself not so trivial. Of course, lissencephalic animals could not be studied because we could not find these landmarks, but why would it mean that they do not have a prefrontal cortex? The reviewer implies that species that we did not study do not have a prefrontal cortex, which makes little sense. Standards in the field of comparative anatomy of the PFC, especially when it implies rodents (lissencephalic also) include cytoarchitectonic and connectivity criteria, but obviously we are not in a position to address it here. We have, however, included references to the seminal work of Angela Roberts and collaborator in the discussion on marmosets prefrontal functions, to reinforce the idea that the functional organization is relatively well conserved across all primates (with or without gyri on their brain) (Dias et al, 1996; Roberts et al, 2007).

Do these primates have significantly more simplistic ways of life than gyrencephalic primates? Marmosets and squirrel monkeys have quite small brains; does it imply that they have not experience the influence of socio-ecological factors on the size of FP, DLPFC, and the rest of the brain?

Again, none of this is relevant here, because we could not draw conclusions on species that we cannot study for methodological reasons. The reviewer seems to believe that an absence of evidence is equivalent to an evidence of absence, but we do not.

The authors state that:

"the strong development of executive functions in species with larger prefrontal cortices is related to an absolute increase in number of neurons, rather than in an increase in the ration between the number of neurons in the PFC vs the rest of the brain".

How does it apply to marmosets and squirrel monkeys?

Again, we do not understand the reviewer’s point, since it is widely admitted that lissencephalic monkeys display both a prefrontal cortex and executive functions (again, see the work of Angela Roberts cited above). Our goal here was certainly not to get into the debate of what is the prefrontal cortex in a handful of laboratory species, but to evaluate the relevance of laboratory based neuro-cognitive concepts for understanding primates in general, and in their natural environment.

References:

García-Cabezas MA, Hacker JL, Zikopoulos B (2022) Homology of neocortical areas in rats and primates based on cortical type analysis: an update of the Hypothesis on the Dual Origin of the Neocortex. Brain structure & function Online ahead of print. doi:doi.org/ 10.1007/s00429-022-02548-0

García-Calero E, Puelles L (2020) Histogenetic radial models as aids to understanding complex brain structures: The amygdalar radial model as a recent example. Front Neuroanat 14:590011. doi:10.3389/fnana.2020.590011

Nieuwenhuys R, Puelles L (2016) Towards a New Neuromorphology. doi:10.1007/978-3-319-25693-1

Puelles L, Alonso A, Garcia-Calero E, Martinez-de-la-Torre M (2019) Concentric ring topology of mammalian cortical sectors and relevance for patterning studies. J Comp Neurol 527 (10):1731-1752. doi:10.1002/cne.24650

Reviewer #3 (Public Review):

This is an interesting manuscript that addresses a longstanding debate in evolutionary biology - whether social or ecological factors are primarily responsible for the evolution of the large human brain. To address this, the authors examine the relationship between the size of two prefrontal regions involved in metacognition and working memory (DLPFC and FP) and socioecological variables across 16 primate species. I recommend major revisions to this manuscript due to: 1) a lack of clarity surrounding model construction; and 2) an inappropriate treatment of the relative importance of different predictors (due to a lack of scaling/normalization of predictor variables prior to analysis). My comments are organized by section below:

We thank the reviewer for the globally positive evaluation and for the constructive remarks. Introduction:

• Well written and thorough, but the questions presented could use restructuring.

Again, we thank the reviewer, and we believe that this is coherent with some of the remarks of reviewer 1. We have extensively revised the introduction, toning down the social vs ecological brain issue to focus more on what is the objective of the work (evaluating the relevance of lab based neuro-cognitive concepts for understanding natural behavior in primates).

Methods:

• It is unclear which combinations of models were compared or why only population density and distance travelled tested appear to have been included.

The details of the model comparison analysis were presented as a table in the supplementary material (#3, details of the model comparison data), but we understand that this was not clear enough. We have provided more explanation both in the main manuscript and in the supplements. All variables were considered a priori; however, we proceeded beforehand to an exploratory analyses which led us to exclude some variables because of their lack of resolution (not enough categories for qualitative variables) or strong cross-correlations with other quantitative variables. There were much more than three variables included in the models but the combination of these 3 (body mass, daily traveled distance and population density) best predicted (had the smallest AIC) the size of the brain regions. We provide additional information about these exploratory analyses in the supplementary material, sections 2 and 3.

• Brain size (vs. body size) should be used as a predictor in the models.

We do not understand the theoretical reason for replacing body size by brain size in the models. Brain size is not a socio-ecological variable. And of course, that would be impossible for modeling brain size itself. Or is it that the reviewer suggests to use brain size as a covariate to evaluate the effects of other variables in the model over and above the effect on brain size? But what is the theoretical basis for this?

• It is not appropriate to compare the impact of different predictors using their coefficients if the variables were not scaled prior to analysis.

We thank the Reviewer for this comment; however, standardized coefficients are not unproblematic because their calculations are based on the estimated standard-deviations of the variables which are likely to be affected by sampling (in effect more than the means). We note that the methods of standardized coefficients have attracted several criticisms in the literature (see the References section in https://en.wikipedia.org/wiki/Standardized_coefficient). Nevertheless, we now provide a table with these coefficients which makes an easy comparison for the present study. We also updated tables 1, 2 and 3 to include standardized beta values.

Reviewer #1 (Recommendations For The Authors):

N/A

Reviewer #2 (Recommendations For The Authors):

Contemporary developmental biology has showed that the brain of all mammals, including primates, develops out of a bauplan (or blueprint) made of several fundamental morphological units that have invariant topological relations across species (Nieuwenhuys and Puelles 2016).

At some point in the discussion the authors acknowledge that:

"Our aim here was clearly not to provide a clear identification of anatomical boundaries across brain regions in individual species, as others have done using much finer neuroanatomical methods. Such a fine neuroanatomical characterization appears impossible to carry on for a sample size of species compatible with PGLS".

I do not think it would be impossible to carry such neuroanatomical characterization. It would take time and effort, but it is feasible. Such characterization, if performed within the framework of contemporary developmental biology, would allow for well-founded definition and delineation of cortical sectors across primate species, including lissencephalic ones, and would allow for meaningful homologies and interspecies comparisons.

We do not see how our work would benefit from developmental biology at that point, because it is concerned with evolution, and these are very distinct biological phenomena. We do not understand the reviewer’s focus on lissencephalic species, because they are not so prevalent across primates, and it is unlikely that adding a couple of lissencephalic species will change much to the conclusions.

Minor points:

Please, format references according to the instructions of the journal.

Ok - done

The authors could use the same color code across Figures 1, 2, and 3.

Ok – done

The authors say that group hunting "only occurs in a few primate species", but it also occurs in wolves, whales, and other mammalian species.

We focus on primates here, these other species are irrelevant. Again, this is beside the point.

Reviewer #3 (Recommendations For The Authors):

My comments are organized by section below:

Introduction:

• Well written and thorough

• The two questions presented towards the end of the intro are not clear and do not guide the structure of the methods/results sections. I believe one it would be more appropriate to ask if: 1) the relative proportions of the FP and DLPFC (relative to ROB) are consistent across primates; and 2) if the relative size of these region is best predicted by social and/ or ecological variables. Then, the results sections could be organized according to these questions (current results section 1 = 1; current results sections 2, 3, 4 = 2.1, 2.2, 2.3)

As explained above, we agree with the reviewer that the introduction was somehow misleading and we have edited it extensively. We do not, however, agree with the reviewer regarding the relative (vs absolute) measure. We have discussed this in our response to reviewer 1 regarding the comparison of regional volumes as proxies for number of neurons. The best predictor of the computing capacity of a brain region is its number of neurons, but there is no reason to believe that this capacity should decrease if the rest of the brain increases, as implied by the relative measure that the reviewer proposes. That debate is probably critical in the field of comparative neuroanatomy, and confronting different perspectives would surely be both interesting and insightful, but we feel that it is beyond the scope of the present article.

Methods:

• While the methods are straightforward and generally well described, it is unclear which combinations of models were compared or why only population density and distance travelled tested appear to have been included (in e.g., Fig SI 3.1) even though many more variables were collected.

We agree that this was not clear enough, and we have tried to improve the description of our model comparison approach, both in the main text and in the supplementary material.

• Why was body mass rather than ROB used as a predictor in the models? The authors should instead/also include analyses using ROB (so the analysis is of FP and DLPFC size relative to brain size). Using body mass confounds the analyses since they will be impacted by differences in brain size relative body size.

Again, we have addressed this issue above. First, body size is a socio-ecological variable (if anything, it especially predicts energetic needs and energy expenditure), but ROB is clearly not. We do not see the theoretical relevance of ROB in a socio-ecological model. Second, from a neurobiological point of view, since within primates the volume of a given brain region is directly related to its number of neurons (again, see work of Herculano-Houzel), which is a good proxy for its computing capacity, we do not see the theoretical reason for considering ROB.

• It is not appropriate to compare the impact of different predictors using their coefficients if the variables were not scaled prior to analysis. The authors need to implement this in their approach to make such claims.

We thank the reviewer again for pointing that out. We have addressed this question above.

• Differences across primates in terms of frontal lobe networks throughout the brain should be acknowledged (e.g., Barrett et al. 2020, J Neurosci).

We have added that reference to the discussion, together with other references showing that the difference between human and non-human primates is significant, but essentially quantitative, rather than qualitative (the building blocks are relatively well conserved, but their relative weight differs a lot). Thank you for pointing it out.

I hope the authors find my comments helpful in revising their manuscript.

And we thank again the reviewer for the helpful and constructive comments.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The diversity of cognitive skills across primates remains both a fascinating and a controversial issue. Recent comparative studies provided conflicting results regarding the contribution of social vs ecological constraints to the evolution of cognition. Here, we used an interdisciplinary approach combining comparative cognitive neurosciences and behavioral ecology. Using brain imaging data from 16 primate species, we measured the size of two prefrontal brain regions, the frontal pole (FP) and the dorso-lateral prefrontal cortex (DLPFC), respectively involved in metacognition and working memory, and examined their relation to a combination of socio-ecological variables. The size of these prefrontal regions, as well as the whole brain, was best explained by three variables: body mass, daily travelled distance (an index of ecological constraints) and population density (an index of social constraint). The strong influence of ecological constraints on FP and DLPFC volumes suggests that both metacognition and working memory are critical for foraging in primates. Interestingly, FP volume was much more sensitive to social constraints than DLPFC volume, in line with laboratory studies showing an implication of FP in complex social interactions. Thus, our data clarifies the relative weight of social vs ecological constraints on the evolution of specific prefrontal brain regions and their associated cognitive operations in primates.

Revised version: The diversity of cognitive skills across primates remains both a fascinating and a controversial issue. Recent comparative studies provided conflicting results regarding the contribution of social vs ecological constraints to the evolution of cognition. Here, we used an interdisciplinary approach combining comparative cognitive neurosciences and behavioral ecology. Using brain imaging data from 16 primate species, we measured the size of two prefrontal brain regions, the frontal pole (FP) and the dorso-lateral prefrontal cortex (DLPFC), respectively involved in metacognition and working memory, and examined their relation to a combination of socio-ecological variables. The size of these prefrontal regions, as well as the whole brain, was best explained by three variables: body mass, daily traveled distance (an index of ecological constraints) and population density (an index of social constraints). The strong influence of ecological constraints on FP and DLPFC volumes suggests that both metacognition and working memory are critical for foraging in primates. Interestingly, FP volume was much more sensitive to social constraints than DLPFC volume, in line with laboratory studies showing an implication of FP in complex social interactions. Thus, our data highlights the relative weight of social vs ecological constraints on the evolution of specific prefrontal brain regions and their associated cognitive operations in primates.

Introduction

Earlier version: Numerous studies have addressed the mechanisms underlying the evolution of cognitive abilities in primates, using brain size as a proxy (Chambers, Heldstab, & O’Hara, 2021; A. R. DeCasien, Barton, & Higham, 2022; Powell, Isler, & Barton, 2017; Carel P. van Schaik, Triki, Bshary, & Heldstab, 2021). One of the major hypotheses, referred to as the ‘social brain’ hypothesis, proposes that social interactions require higher cognitive skills such that the complexity of social interactions played a central role in the increase in brain size during primate evolution (Dunbar, 1998). Another major theory, referred to as the ‘ecological brain’ hypothesis, proposes that the increase in brain size and corresponding cognitive skills was driven by the need to build mental representations in order to forage efficiently (Milton, 1981, 1993). Even if both of these hypotheses are supported by experimental evidence, they are often opposed, with the idea that one of them should have significantly more weight on brain evolution compared to the other (e.g. (Alex R. DeCasien, Williams, & Higham, 2017)). Critically, this apparent opposition between the ‘social brain’ and the ‘ecological brain’ is based on the assumption that the whole brain volume is a good proxy for overall cognitive abilities (Deaner, Isler, Burkart, & van Schaik, 2007; Falk & Gibson, 2001). In this framework, cognitive skills could allow to solve both ecological and social challenges, but one of these two factors (social or ecological) would have had more weight on the evolution of these skills across species.

Revised version: Numerous studies have addressed the mechanisms underlying the evolution of cognitive abilities in primates, using brain size as a proxy (Powell, Isler et al. 2017, Chambers, Heldstab et al. 2021, van Schaik, Triki et al. 2021, DeCasien, Barton et al. 2022). One of the major hypotheses, referred to as the ‘social brain’ hypothesis, proposes that social interactions require higher cognitive skills such that the complexity of social interactions played a central role in the increase in brain size during primate evolution (Dunbar 1998). Another major theory, referred to as the ‘ecological brain’ hypothesis, proposes that the increase in brain size and corresponding cognitive skills was driven by the need to build mental representations in order to forage efficiently (Milton 1981, Milton 1993). Critically, both of these theories are based on the assumption that the whole brain volume is a good proxy for overall cognitive abilities, and that these global abilities can benefit social interactions or foraging, across species in an evolutionary framework. But clearly, the level of precision of such neuro-cognitive scenario remains limited, given the precision with which neuroscience and psychology have characterized the brain functional anatomy and cognition, respectively.

Earlier version: Nevertheless, this assumption remains poorly specified from a psychological and a neuro-physiological perspective. Indeed, from a neurobiological perspective, the idea that the social and ecological brain hypotheses should be in opposition appears counterintuitive, since recent studies have shown that, for the most part, the brain regions involved in social interactions vs. foraging were different. One one hand, imaging studies in humans have identified a set of specific brain regions (e.g., the temporo-parietal junction, TPJ) involved in social interactions (Stanley & Adolphs, 2013). The network of brain regions involved in social interactions is often referred to as the social brain, and recent comparative studies have shown that it exists in other primate species including macaques (Rushworth, Mars, & Sallet, 2013; Jérôme Sallet et al., 2011; Testard et al., 2022). On the other hand, foraging, which involves a complex spatio-temporal representation of food availability and values, would involve a distinct set of brain regions, including the hippocampus and the ventromedial prefrontal cortex (Lin, Horner, Bisby, & Burgess, 2015; Louail, Gilissen, Prat, Garcia, & Bouret, 2019; Vikbladh et al., 2019; Zuberbühler & Janmaat, 2010). However, the brain systems potentially involved in social cognition and foraging are probably not completely independent: several regions of the frontal lobe seem to be involved both in foraging and social interactions (Barbey et al., 2014; Gallagher & Frith, 2003; Kolling, Behrens, Mars, & Rushworth, 2012; Mansouri, Koechlin, Rosa, & Buckley, 2017; Rudebeck, Walton, Smyth, Bannerman, & Rushworth, 2006; Yoshida, Saito, Iriki, & Isoda, 2012). This partial overlap between the so called ‘social brain’ and the so called ‘foraging brain’ is in line with the idea that some cognitive operations are generic enough to be involved in multiple contexts (Garcia, Bouret, Druelle, & Prat, 2021). This is typically the case for executive functions, which allow to organize behavior over space and time to reach a goal, and rely upon several regions of the prefrontal cortex (Fuster, 2008; Luria, 1973). But the extent to which these functional networks identified in laboratory conditions apply to behavior in the wild often remains very speculative. Indeed,, even if the mapping between cognitive functions and specific brain regions has been a core issue for cognitive neurosciences for several decades, it has mostly been studied in laboratory conditions and in clinical settings (Passingham, Passingham, & Wise, 2012; Yechiam, Busemeyer, Stoof, & Bechara, 2005). But to what extent are these cognitive operations and their neural substrate identified in laboratory settings related to natural socio-ecological processes, and to the socio-ecological variability across species, based on their evolutionary history?

Revised version: Imaging studies in humans have identified a network of brain regions involved in social interactions, and this network is often referred to as the social brain (Stanley and Adolphs 2013). Recent comparative studies have shown that it also existed in macaques (Sallet, Mars et al. 2011, Rushworth, Mars et al. 2013, Testard, Brent et al. 2022). Along the same lines, laboratory studies have also proposed that a set of specific brain regions (e.g. cingulate cortex) was underlying foraging in humans and macaques (eg (Hayden, Pearson et al. 2011, Kolling, Behrens et al. 2012). These studies, however, have only been conducted in a handful of species and in very artificial conditions, such that the relation between these brain regions and natural behavior or evolution remains elusive. In terms of the cognitive processes involved, ‘social interactions’ and ‘foraging’ remain poorly specific and at least in primates, they probably involve myriads of more elementary cognitive operations. For example, social interactions involve some form of categorization of facial expressions and metacognition, and the brain regions associated with these elementary cognitive operations (e.g. frontal pole or temporo-parietal junction) are usually part of the ‘social brain’, i.e. the set of brain regions associated with social interactions (Gallagher and Frith 2003, Rushworth, Mars et al. 2013, Devaine, San-Galli et al. 2017, de Gelder and Poyo Solanas 2021, de Gelder 2023, Deen, Schwiedrzik et al. 2023). Along the same lines, foraging is thought to require mental spatio-temporal representations of food availability and values, and these elementary cognitive operations rely upon a distinct set of brain regions, including the hippocampus and the ventromedial prefrontal cortex (Zuberbühler and Janmaat 2010, Lin, Horner et al. 2015, Louail, Gilissen et al. 2019, Vikbladh, Meager et al. 2019). Finally, the brain systems potentially involved in social cognition and foraging partially overlap, especially in the frontal lobes (Gallagher and Frith 2003, Rudebeck, Walton et al. 2006, Kolling, Behrens et al. 2012, Yoshida, Saito et al. 2012, Barbey, Colom et al. 2014, Mansouri, Koechlin et al. 2017). This partial overlap between the networks referred to as the ‘social brain’ and the ‘ecological brain’ is in line with the idea that some elementary cognitive operations are generic enough to be involved in multiple contexts (Garcia, Bouret et al. 2021, Shultz and Dunbar 2022). This is typically the case for executive functions (e.g. working memory), which allow to organize behavior over space and time to reach a goal, and rely upon several regions of the prefrontal cortex (Luria 1973, Fuster 2008).

Added in the revised version: Given the complexity encountered in natural environments, it is hazardous to extrapolate these findings in artificial laboratory conditions and in a handful of species to the decision-making processes in both social and food dimensions occurring in wild primates, for which the diversity of social interactions (e.g. number of individuals an animal can remember, degree of social awareness i.e. knowledge and representation of the dominance hierarchy, kinship relations or friendship associations) and foraging behaviors (e.g. food processing techniques, knowledge of harvesting schedules, extractive foraging) require much more cognitive flexibility. One can wonder to what extent these neuro-cognitive concepts (i.e. elementary cognitive operation and their underlying cerebral substrate) can be used as theoretical bricks to understand the diversity of natural behaviors reported in the wild and their evolution across species.

Added in the revised version: In terms of evolution, the functional heterogeneity of distinct brain regions is captured by the notion of ‘mosaic brain’, where distinct brain regions could show a specific relation with various socio-ecological challenges, and therefore have relatively separate evolutionary trajectories (Barton and Harvey 2000, DeCasien and Higham 2019). But these relations between specific brain regions and socio-ecological variables across species provide little insight into the cognitive processes that could be at play. Thus, it remains challenging to bridge the gap between 1) laboratory studies relating brain regions to elementary cognitive operations in individual species and 2) studies relating brain regions to socio-ecological challenges through an evolutionary scenario (i.e. across numerous species).

Removed in the revised version: Of course, the precision with which these cognitive operations and their neurobiological substrate have been defined in laboratory conditions is quite challenging to achieve in wild conditions, and even more so when they must be compared across species (but see Janmaat et al, 2021). In terms of evolution, the functional heterogeneity of distinct brain regions is captured by the notion of ‘mosaic brain’, where distinct brain regions could show a specific relation with various socio-ecological challenges, and therefore have relatively separate evolutionary trajectories (Barton & Harvey, 2000; Alex R. DeCasien & Higham, 2019b). What remains challenging, however, is to bridge the gap between brain structures and functions associated with specific cognitive operations in laboratory conditions and brain functions in natural conditions, when facing real socio-ecological challenges. Again, in laboratory conditions, cognitive operations are captured very precisely but their relation to socio-ecological challenges remains, at best, speculative. In order to bridge this gap, it seems essential to derive predictions where a set of well identified cognitive operations and their associated brain regions studied in laboratory conditions could be involved in more natural conditions as assessed in wild animals evolving in their natural environments (Barks et al., 2014)(A. R. DeCasien et al., 2022; Janmaat et al., 2021; Louail et al., 2019).

Earlier version: For this study, we focused on two well-known cognitive operations, both considered as executive functions, and involving distinct regions of the prefrontal cortex: metacognition which involves the frontal pole (FP) and working memory involving the dorso-lateral prefrontal cortex (DLPFC). First, metacognition might be critical both for social interactions by supporting theory of mind (the ability of an individual to conceptualize others’ states of mind) and for foraging by enabling complex planning (Devaine et al., 2017; Fleming & Dolan, 2012; Frith, 2007; Mansouri et al., 2017). For example, group hunting, which involves both social and foraging functions, only occurs in a few primate species where metacognition is thought to be particularly developed (Boesch, 1994; Conard, Serangeli, Bigga, & Rots, 2020; Garcia et al., 2021; C. P. van Schaik, 2016). Secondly, given the very clear role of DLPFC in working memory and planning, we could speculate that it is critically involved in foraging, but its role in complex social interactions is less clear: on one hand, it is not often associated with social interactions in laboratory studies, but on the other hand working memory and planning could readily be involved in complex social interactions in more natural conditions (Frith, 2007; Fuster, 2008; Passingham et al., 2012; J. Sallet et al., 2013).

Revised version: In order to bridge this gap, we designed this study to derive and test predictions where a set of well-identified elementary cognitive operations and their associated brain regions (as studied in laboratory conditions) could be involved in more natural functions (as assessed in wild animals living in their natural environments). We focused on two well-known cognitive operations, both considered as executive functions, and involving distinct regions of the prefrontal cortex: metacognition, which involves the frontal pole (FP), and working memory, which involves the dorso-lateral prefrontal cortex (DLPFC), as demonstrated by extensive work in humans and macaques (Fuster 2008, Passingham, Passingham et al. 2012, Mansouri, Koechlin et al. 2017). Based on the literature, we hypothesized that FP and metacognition would be associated both with social interactions (by supporting theory of mind, the ability of an individual to conceptualize others’ states of mind) and with foraging (by enabling complex planning) (Frith 2007, Fleming and Dolan 2012, Devaine, San-Galli et al. 2017, Mansouri, Koechlin et al. 2017). For example, group hunting, which involves both social and foraging functions, only occurs in a few primate species where metacognition is thought to be particularly developed (Boesch 1994, van Schaik 2016, Conard, Serangeli et al. 2020, Garcia, Bouret et al. 2021). Concerning the DLPFC, given its very clear implication in working memory and planning, we speculated that it would be critically involved in foraging. Its potential relation with social interactions, however, was less clear: on one hand, DLPFC is rarely associated with social interactions in laboratory studies, but on the other hand working memory and planning could readily be involved in complex social interactions in more natural conditions (Frith 2007, Fuster 2008, Passingham, Passingham et al. 2012, Sallet, Mars et al. 2013).

Earlier version: In order to bridge the gap between these brain regions (and their functions) defined in laboratory conditions and potential neuro-cognitive operations that evolved across primates to adapt to their specific socio-ecological challenges in the wild, we addressed two central questions: 1) to what extent does this functional mapping of prefrontal regions generalize across primates, and 2) to what extent does the functional mapping between brain regions and specific executive functions observed in laboratory conditions translate to natural conditions?

Revised version: To evaluate the relevance of laboratory-based neuro-cognitive concepts for understanding which cognitive operations could be mobilized to overcome specific socio-ecological challenges in the wild and their variability across species, we addressed two central questions: 1) to what extent does this functional mapping of prefrontal regions generalize across primates, and 2) to what extent does the functional mapping between brain regions and specific executive functions observed in laboratory relate to socio-ecological variables measured in natural conditions?

Earlier version: To address these questions, we measured the volume of the two prefrontal regions of interest (FP and DLPFC) in 16 primate species and evaluated the influence of key socio-ecological factors on the size of these regions. Based on the known positive relation between the size of a given brain region (or the number of corresponding neurons), and the relative importance of its associated function, both within and across species, we assume that the relative size of each brain region can be taken as a proxy for its associated function (metacognition for FP and working memory for DLPFC) (Barks et al., 2014; Barton & Harvey, 2000; Alex R. DeCasien & Higham, 2019a; Ferrucci et al., 2022; Herculano-Houzel, 2017; Herculano-Houzel, Kaas, & de Oliveira-Souza, 2016; Louail et al., 2019; Maguire et al., 2000; Jérôme Sallet et al., 2011; Testard et al., 2022). These two regions were chosen and identified based on functional maps and reliable macroscopic landmarks, rather than cytoarchitectonic criteria, to maximize the reliability of the comparative approach (see also (Louail et al., 2019). Anatomical studies have shown that these regions and landmarks were relatively well conserved between humans and macaques, both in terms of cytoarchitectonics and connectivity profiles (Amiez et al., 2019; Michael Petrides, Tomaiuolo, Yeterian, & Pandya, 2012; J. Sallet et al., 2013). Finally, the cognitive functions in which they are involved appear relatively well conserved between macaques and humans (Fuster, 2008; Mansouri et al., 2017). In addition to these two regions, we also used measures of whole brain volume as a reference to facilitate comparison with other studies that only use whole brain measures. The 16 primate species were chosen to cover a wide range of phylogenetic distances coupled to a diversity of socio-ecological variables. We compared the influence of several combinations of these socio-ecological variables on the variability in size of each brain region of interest, and interpreted it in the light of its known function in laboratory conditions. Thus, we could evaluate the strength of the relation between specific regions of the prefrontal cortex and specific socio-ecological variables across species, as observed in the wild, thereby complementing laboratory studies and bridging the gap between cognitive neurosciences, behavioral ecology and primate evolution.

Revised version: To address these questions, we measured the volume of the two prefrontal regions of interest (FP and DLPFC) in 16 primate species and evaluated the influence of key socio-ecological factors on the size of these regions. Based on the known positive relation between the size of a given brain region (or the number of corresponding neurons), and the relative importance of its associated function, both within and across species, we assume that the size of each brain region can be taken as a proxy for the weight of its associated function on behavior (Barton and Harvey 2000, Maguire, Gadian et al. 2000, Sallet, Mars et al. 2011, Barks, Calhoun et al. 2014, Herculano-Houzel, Kaas et al. 2016, Herculano-Houzel 2017, DeCasien and Higham 2019, Louail, Gilissen et al. 2019, Ferrucci, Nougaret et al. 2022, Testard, Brent et al. 2022). These two regions were chosen and identified based on functional maps and reliable macroscopic landmarks, rather than cytoarchitectonic criteria, to maximize the reliability of the comparative approach (see also (Louail, Gilissen et al. 2019). Anatomical studies have shown that these regions and landmarks were relatively well conserved between humans and macaques, both in terms of cytoarchitectonics and connectivity profiles (Petrides, Tomaiuolo et al. 2012, Sallet, Mars et al. 2013, Amiez, Sallet et al. 2019). Finally, the cognitive functions in which they are involved appear relatively well conserved between macaques and humans (Fuster 2008, Mansouri, Koechlin et al. 2017). In addition to these two regions, we also used measures of whole brain volume as a reference to facilitate comparison with other studies that only use whole brain measures. The 16 primate species were chosen to cover a wide range of phylogenetic distances coupled to a diversity of socio-ecological variables. We compared the influence of several combinations of these socio-ecological variables on the variability in size of each brain region of interest, and interpreted it in the light of its known function in laboratory conditions. Thus, we could evaluate the strength of the relation between specific regions of the prefrontal cortex and specific socio-ecological variables across species, as observed in the wild, thereby complementing laboratory studies and bridging the gap between cognitive neurosciences, behavioral ecology and primate evolution.

Material & Methods

Sample

Earlier version: Thirty one brain magnetic resonance (MRI) 3D reconstructions from 16 primate species (Ateles fusciceps, n=1; Cebus capucinus, n=1; Cercopithecus mitis, n= 1; Gorilla gorilla, n=5; Gorilla beringei, n=1; Homo sapiens, n=4; Hylobates lar, n=1; Lagothrix lagotricha, n=1; Lophocebus albigena, n=1; Macaca fascicularis, n=1; Macaca fuscata, n=4; Macaca mulatta, n=2; Pan troglodytes, n=5; Pan paniscus, n=1; Papio papio, n=1, Pongo pygmaeus, n=1) were used in this study. Japanese macaques (M. fuscata) and rhesus macaques (M. mulatta) were captive animals scanned at the National Institutes for Quantum and Radiological Science and Technology (Chiba, Japan) and at Brain and Spine Institute (Paris, France), respectively. P. troglodytes and G. gorilla brains came from the Museum national d’Histoire naturelle (Paris, France). They had been collected between 1920 and 1970 and subsequently preserved in formalin solution. The P. paniscus and G. beringei brains came from the Royal Museum for Central Africa (RMCA) (Tervuren, Belgium) and the Royal Belgian Institute of Natural Sciences (Bruxelles, Belgium). All the Pan and Gorilla specimens have been scanned at University of Leuven (KUL). The brain scans A. fusciceps, C. mitis, and L. lagotricha were obtained from the Primate Brain Bank, NIN Utrecht University. Finally, the remaining species (C. capucinus, H. lar, L. albigena, P. pygmaeus, M. fascicularis, P. papio and one specimen of H. sapiens) came from the brain catalogue website (https://braincatalogue.org), which gathers scans of specimens from the collections of the Museum national d’Histoire naturelle (Paris). The three other H. sapiens brain scans were obtained from the Allen Institute (online brain atlas). All specimens were sexually mature at the time of scanning. The sexes of individuals were mostly unknown. Moreover, some specimens in the sample came from captivity. Thus, we neglected the effects of captivity and sex on brain/endocranium measurements, which were both shown to be very small compared to inter-species differences (Isler & van Schaik, 2012).

Revised version: Thirty one brain magnetic resonance (MRI) 3D reconstructions from 16 primate species (Ateles fusciceps, n=1; Cebus capucinus, n=1; Cercopithecus mitis, n= 1; Gorilla gorilla, n=5; Gorilla beringei, n=1; Homo sapiens, n=4; Hylobates lar, n=1; Lagothrix lagotricha, n=1; Lophocebus albigena, n=1; Macaca fascicularis, n=1; Macaca fuscata, n=4; Macaca mulatta, n=2; Pan troglodytes, n=5; Pan paniscus, n=1; Papio papio, n=1, Pongo pygmaeus, n=1) were used in this study. Ja-panese macaques (M. fuscata) and rhesus macaques (M. mulatta) were captive animals scanned at the National Institutes for Quantum and Radiological Science and Technology (Chiba, Japan) and at Brain and Spine Institute (Paris, France), respectively. P. troglodytes and G. gorilla brains came from the Muséum national d’Histoire naturelle (Paris, France). They had been collected between 1920 and 1970 and subsequently preserved in formalin solution. The P. paniscus and G. beringei brains came from the Royal Museum for Central Africa (RMCA) (Tervuren, Belgium) and the Royal Belgian Institute of Natural Sciences (Bruxelles, Belgium). All the Pan and Gorilla specimens have been scanned at University of Leuven (KUL). The brain scans A. fusciceps, C. mitis, and L. lagotricha were obtained from the Primate Brain Bank, NIN Utrecht University. Finally, the remaining species (C. capucinus, H. lar, L. albigena, P. pygmaeus, M. fascicularis, P. papio and one specimen of H. sapiens) came from the brain catalogue website (https://braincatalogue.org), which gathers scans of specimens from the collections of the Muséum national d’Histoire naturelle (Paris). The three other H. sapiens brain scans were obtained from the Allen Institute (online brain atlas). All specimens were sexually mature at the time of scanning. The sexes of individuals were mostly unknown. Moreover, some specimens in the sample came from captivity. Thus, we neglected the effects of captivity and sex on brain/endocranium measurements, which were both shown to be very small compared to inter-species differences (Isler and van Schaik 2012).

Processing of brain MRI and measurements

Brain measurements (visualization, segmentation and quantification of brain tissues volumes) were processed using Avizo v9.0 software. The whole brain volume was measured in order to facilitate comparisons with the literature. The cerebellum was excluded from all brain measurements, because it was missing on some of the MRI scans (gorilla brains). Whole brain segmentation was performed using the semi-automated tool in Avizo that enables to select a material or structure according to a specific gray-level threshold. It was however necessary to bring some manual corrections, for example when the brain had a similar gray-level than an adjacent tissue. Segmentations of the frontal pole (FP) and the dorso-lateral prefrontal cortex (DLPFC) were carried out manually with the brush tool, slice by slice of the MRI scan.

Earlier version: The frontal pole is the most rostral part of the prefrontal cortex. Cytoarchitectonic studies indicate that it strongly overlaps with Brodmann area 10 (BA 10) (Ongür, Ferry, & Price, 2003; Ramnani & Owen, 2004; Semendeferi, Armstrong, Schleicher, Zilles, & Van Hoesen, 2001; Tsujimoto, Genovesio, & Wise, 2011). Besides cytoarchitectonic landmarks, the frontal pole can also be identified based on connectivity patterns: it receives projections from the temporal superior cortex (M. Petrides, 2005) and has connections with the superior temporal sulcus (J. Sallet et al., 2013). We also used the probabilistic maps of the frontal pole proposed by John et al. (John et al., 2007) and Bludau et al. (Bludau et al., 2014). We delimited the frontal pole according to different criteria: it should match the functional anatomy for known species (macaques and humans, essentially) and be reliable enough to be applied to other species using macroscopic neuroanatomical landmarks. We integrated these criteria and the data from the literature on brain atlases of rhesus macaques and humans (Borden, Stefan, & Forseen, 2015; Saleem & Logothetis, 2007) to define visible limits of the frontal pole, as shown in Figure 1. The anterior limit of the frontal pole was defined as the anterior limit of the brain. The cingulate sulcus represented the posterior limit. The ventral limit was set as the dorsal limit of the gyrus rectus. Finally, the dorsal limit was defined as the fundus of the superior frontal sulcus in humans and apes, or the principal sulcus in monkeys.

Revised version: The frontal pole is the most rostral part of the prefrontal cortex. Cytoarchitectonic studies indicate that it strongly overlaps with Brodmann area 10 (BA 10) (Semendeferi, Armstrong et al. 2001, Ongür, Ferry et al. 2003, Ramnani and Owen 2004, Tsujimoto, Genovesio et al. 2011). Besides cytoarchitectonic landmarks, the frontal pole can also be identified based on connectivity patterns: it receives projections from the temporal superior cortex (Petrides 2005) and has connections with the superior temporal sulcus (Sallet, Mars et al. 2013). We also used the probabilistic maps of the frontal pole proposed by John et al. (John, Yashavantha et al. 2007) and Bludau et al. (Bludau, Eickhoff et al. 2014). We delimited the frontal pole according to different criteria: it should match the functional anatomy for known species (macaques and humans, essentially) and be reliable enough to be applied to other species using macroscopic neuroanatomical landmarks. We integrated these criteria and the data from the literature on brain atlases of rhesus macaques and humans (Saleem and Logothetis 2007, Borden, Stefan et al. 2015) to define visible limits of the frontal pole, as shown in Figure 1. The anterior limit of the frontal pole was defined as the anterior limit of the brain. The cingulate sulcus represented the posterior limit. The ventral limit was set as the dorsal limit of the gyrus rectus. Finally, the dorsal limit was defined as the fundus of the superior frontal sulcus in humans and apes, or the principal sulcus in monkeys. We chose this very conservative dorsal limit in monkeys to avoid any overlap with what could be considered as part of the DLPFC (typically, BA9). However, we are aware that several studies in macaques suggest that the frontal pole in monkeys might extend more dorsally (Tsujimoto, Genovesio et al. 2011, Sallet, Mars et al. 2013). Thus, we also conducted a series of measures with a less conservative dorsal limit for the frontal pole (see supplementary material section 4).

Boundaries of the frontal pole. From left to right: sagittal view, coronal view for great apes and humans, coronal view for monkeys. Frontal pole is in green. Abbreviations: cs: cingulate sulcus; cc: corpus callosum; rs: rostral sulcus; mos: medial-orbital sulcus; frss: superior frontal sulcus; ps: principal sulcus; D: dorsal; V: ventral; M: medial; L: lateral; A: anterior; P: posterior.

Earlier version: Similarly, the dorso-lateral prefrontal cortex was measured by combining functional and anatomical data from the literature to identify reliable macroscopic landmarks (Levy & Goldman-Rakic, 2000; Passingham et al., 2012; Passingham & Sakai, 2004; Michael Petrides et al., 2012; J. Sallet et al., 2013). Given the major difference in sulcal organization between monkeys and greats apes, we used different landmarks, shown on figures 2a and 2b. DLPFC comprises portions of middle frontal gyrus and superior frontal gyrus in great apes and lies in and around the principal sulcus in macaques. Then, the ventral limit was set as the fundus of the principal sulcus for monkeys, and the frontal inferior sulcus for apes and humans. The anterior limit of the DLPFC was defined as the posterior limit of the frontal pole, which was the cingulate sulcus. For apes and humans, the posterior limit was defined as the precentral sulcus, whereas in monkeys it was defined as the end of the arcuate sulcus. Finally, the medial limit was designated as the inter-hemispheric sulcus.

Revised version: Similarly, the dorso-lateral prefrontal cortex was measured by combining functional and anatomical data from the literature to identify reliable macroscopic landmarks (Levy and Goldman-Rakic 2000, Passingham and Sakai 2004, Passingham, Passingham et al. 2012, Petrides, Tomaiuolo et al. 2012, Sallet, Mars et al. 2013). Given the major difference in sulcal organization between monkeys and greats apes, we used different landmarks, shown on figures 2a and 2b. DLPFC comprises portions of middle frontal gyrus and superior frontal gyrus in great apes and lies in and around the principal sulcus in macaques. Then, the ventral limit was set as the fundus of the principal sulcus for monkeys, and the frontal inferior sulcus for apes and humans. The anterior limit of the DLPFC was defined as the posterior limit of the frontal pole, which was the cingulate sulcus. For apes and humans, the posterior limit was defined as the precentral sulcus, whereas in monkeys it was defined as the end of the arcuate sulcus. Finally, the medial limit was designated as the inter-hemispheric sulcus.

Boundaries of the dorso-lateral prefrontal cortex

a)for great apes and humans. From left to right: sagittal view, external view, coronal view. Abbreviations: cs: cingulate sulcus; cc: corpsus callosum; frsi: frontal inferior sulcus; frss: frontal superior sulcus; cent.s: central sulcus; prec.s: precentral sulcus; A: anterior: P: posterior; D: dorsal; V: ventral; L: lateral; M: medial.

b) for monkeys. From left to right: sagittal view, external view, coronal view. Abbreviations: cs: cingulate sulcus; cc: corpsus callosum; ps: principal sulcus; as: arcuate sulcus; A: anterior: P: posterior; D: dorsal; V: ventral; L: lateral; M: medial.

Socio-ecological and phylogenetic data

Earlier version: Eleven socio-ecological variables were selected for the analyses, gathered in different categories: body condition (body mass), diet (dietary quality index and tool use), movements and ranging behavior (daily traveled distance), social parameters (group size, population density, social system), and variables related to reproduction and life-history traits (mating system, mate guarding, seasonal breeding, and weaning age). Each variable was assessed based on the literature on wild populations, whenever possible, which was the case in a vast majority of cases (and otherwise specified, see below). We verified that these variables showed minimal correlation (see supplementary data section 2).

Revised version: Eleven socio-ecological variables were selected for the analyses, gathered in different categories: body condition (body mass), diet (dietary quality index and tool use), movements and ranging behavior (daily traveled distance), social parameters (group size, population density, social system), and variables related to reproduction and life-history traits (mating system, mate guarding, seasonal breeding, and weaning age). Each variable was assessed based on the literature on wild populations, whenever possible, which was the case in a vast majority of cases (and otherwise specified, see below). We verified that these variables showed minimal correlation (see supplementary material section 2).

Earlier version: The dietary quality index (DQI) was used to characterize the richness of the dietary spectrum and was calculated from the formula DQI = 1s + 2r + 3,5a, where s is the percentage of plant structural parts in the diet, r the percentage of plant reproductive parts and a the percentage of animal preys (Sailer, Gaulin, Boster, & Kurland, 1985). Thereby, a low index (around 100) characterizes folivorous diets, while a high index (around 200) characterizes more diversified diets (including animals and fruits). Tool use represented the occurrence and complexity of using objects in feeding contexts. We took the definition of St Amant and Horton (St Amant & Horton, 2008) which states that “tool use is the exertion of control over a freely manipulable external object (the tool) with the goal of (1) altering the physical properties of another object, substance, surface or medium (the target, which may be the tool user or another organism) via a dynamic mechanical interaction, or (2) mediating the flow of information between the tool user and the environment or other organisms in the environment.” It includes notions of control over an object and goal. The presence or absence of tool use in feeding contexts for species in the wild was assessed from literature. Daily traveled distance was expressed in kilometers. Regarding social parameters, group size was defined as the mean social group size of a species. We also used population density (number of individuals per km2) in order to bring spatial precisions over the group size variable. Group size and population density data were collected and compiled from several primary and secondary sources (see supplemantary material, SI section 1). For group size and population density of Homo sapiens, we took an average between industrialized societies and hunter-gatherers societies. Social system was defined using the four-way categorization scheme typically used in primate studies and included solitary, pair-living, polygyny and polygynandry (see also (Alex R. DeCasien et al., 2017). Mating system categories included spatial polygyny (among solitary species, agonistically powerful males defend mating access to several females), monogamy (one male is socially bonded to one breeding female), polyandry (one female is simultaneously bonded to multiple males), harem polygyny (one male is simultaneously bonded to multiple breeding females) and polygynandry (multiple males and multiple females breed within the same group, but no lasting bonds are formed) (Clutton-Brock, 1989). Mate guarding is defined as a female monopolization over an extended period of time to secure paternity (Manson, 1997) and is characterized by male attempts to associate and copulate with a female during the presumptive fertile period (Dixson & Oxford University Press., 2012; Manson, 1997). This categorical variable was divided into two categories: species using mate guarding and those that do not usually show this behavior. Finally, in order to account for differences in lifespan between study species, the weaning age was calculated as a percentage of maximum lifespan. For some missing values from wild studies, data were taken from studies in captivity and compared to data related to close species (in the wild and in captivity). This method was applied for the weaning age and the dietary quality index of Papio papio, with other species of baboons. Moreover, the absence of tool use for Lophocebus albigena was inferred from the absence of published papers or other forms of communication on this subject.

Revised version: The dietary quality index (DQI) was used to characterize the richness of the dietary spectrum and was calculated from the formula DQI = 1s + 2r + 3,5a, where s is the percentage of plant structural parts in the diet, r the percentage of plant reproductive parts and a the percentage of animal preys (Sailer, Gaulin et al. 1985). Thereby, a low index (around 100) characterizes folivorous diets, while a high index (around 200) characterizes more diversified diets (including animals and fruits). Tool use represented the occurrence and complexity of using objects in feeding contexts. We took the definition of St Amant and Horton (St Amant and Horton 2008) which states that ‘tool use is the exertion of control over a freely manipulable external object (the tool) with the goal of (1) altering the physical properties of another object, substance, surface or medium (the target, which may be the tool user or another organism) via a dynamic mechanical interaction, or (2) mediating the flow of information between the tool user and the environment or other organisms in the environment.’ It includes notions of control over an object and goal. The presence or absence of tool use in feeding contexts for species in the wild was assessed from literature. Daily traveled distance was expressed in kilometers. Regarding social parameters, group size was defined as the mean social group size of a species. We also used population density (number of individuals per km2) in order to bring spatial precisions over the group size variable. Group size and population density data were collected and compiled from several primary and secondary sources (see supplementary material section 1). For group size and population density of Homo sapiens, we took an average between industrialized societies and hunter-gatherers societies. Social system was defined using the four-way categorization scheme typically used in primate studies and included solitary, pair-living, polygyny and polygynandry (see also (DeCasien, Williams et al. 2017). Mating system categories included spatial polygyny (among solitary species, agonistically powerful males defend mating access to several females), monogamy (one male is socially bonded to one breeding female), polyandry (one female is simultaneously bonded to multiple males), harem polygyny (one male is simultaneously bonded to multiple breeding females) and polygynandry (multiple males and multiple females breed within the same group, but no lasting bonds are formed) (Clutton-Brock 1989). Mate guarding is defined as a female monopolization over an extended period of time to secure paternity (Manson 1997) and is characterized by male attempts to associate and copulate with a female during the presumptive fertile period (Manson 1997, Dixson and Oxford University Press. 2012). This categorical variable was divided into two categories: species using mate guarding and those that do not usually show this behavior. Finally, in order to account for differences in lifespan between study species, the weaning age was calculated as a percentage of maximum lifespan. For some missing values from wild studies, data were taken from studies in captivity and compared to data related to close species (in the wild and in captivity). This method was applied for the weaning age and the dietary quality index of Papio papio, with other species of baboons. Moreover, the absence of tool use for Lophocebus albigena was inferred from the absence of published papers or other forms of communication on this subject.

Earlier version: The phylogenetic tree was obtained from the 10ktrees website (https://10ktrees.nunn-lab.org/Primates/downloadTrees.php, version 3). This version (Arnold, Matthews, & Nunn, 2010) provides a Bayesian inference of the primate phylogeny based on collected data for eleven mitochondrial and six autosomal genes from GenBank across 301 primate species.

Revised version: The phylogenetic tree was obtained from the 10ktrees website (https://10ktrees.nunn-lab.org/Pri-mates/downloadTrees.php, version 3). This version (Arnold, Matthews et al. 2010) provides a Bayesian inference of the primate phylogeny based on collected data for eleven mitochondrial and six autosomal genes from GenBank across 301 primate species.

Statistical analysis

We used a phylogenetic generalized least squares (PGLS) approach to evaluate the joint influence of socio-ecological variables on the neuro-anatomical variability across species (Grafen 1989). This approach allowed us to take into account the phylogenetic relation across species when evaluating the influence of socio-ecological variables on their neuroanatomy.

Removed in the revised version: To identify the combination of socio-ecological variables that best predicted the size of a given brain region, given the phylogenetic relations across species, we fitted neuro-anatomical data with several PGLS models, each reflecting a specific combination of socio-ecological variables.

Earlier version: Brain measurements and body mass were log10-transformed before analyses. For the same model several correlation structures were used: standard Brownian, Pagel’s (Pagel, 1999), and OU-based (Martins & Hansen, 1997) and were compared based on the smallest AIC values (Akaike, 1974). For a given correlation structure, the effects of the social and ecological variables were assessed with likelihood-ratio tests. All analyses were performed with ape (Paradis & Schliep, 2019). Because of the relatively small number of species in our sample, we assessed the reliability of the inferred models with a “leave-one-out” procedure: we removed one species from the data and the phylogenetic tree and re-fitted the model selected previously. This was repeated for each species.

Revised version: To identify the combination of socio-ecological variables that best predicted the size of a given brain region, given the phylogenetic relations across species, we fitted neuro-anatomical data with several PGLS models, each reflecting a specific combination of socio-ecological variables. Brain measurements and body mass were log10-transformed before analyses. For the same model several correlation structures were used: standard Brownian, Pagel’s (Pagel 1999), and OU-based (Martins and Hansen 1997) and were compared based on the smallest AIC values (Akaike 1974). For a given correlation structure, the effects of the social and ecological variables were assessed with likelihood-ratio tests. All analyses were performed with ape (Paradis and Schliep 2019).

Added in the revised version: We also calculated standardized coefficients, to facilitate comparison. Standardized coefficients were calculated by the product of the coefficient estimated by PGLS with the ratio of the standarddeviation of the predictor on the standard-deviation of the response.

Added in the revised version: Because of the relatively small number of species in our sample, we assessed the reliability of the inferred models with a ‘leave-one-out’ procedure: we removed one species from the data and the phylogenetic tree and re-fitted the model selected previously. This was repeated for each species. The details of this procedure are provided in supplementary material, section 6.

Results

Neuroanatomical measures

The average size of the regions of interest (whole brain, FP and DLPFC) are shown on figure 3. As expected, all these measures were highly correlated (all r=0.99, Pearson correlation). Not only the volumes of DLPFC and FP, which are two neighboring regions of the prefrontal cortex, but also each of these regions and the whole brain.

Average values of the 3 regions of interest. Each line provides the cumulated volumes of the dorso-lateral prefrontal cortex (DLPFC, green), the frontal pole (FP, grey) and the rest of the brain (ROB, blue), such that the size of each bar represents the whole brain volume. Note: for Homo sapiens, the bar has been truncated, since its value was out of scale with the other species.

Influence of socio-ecological variables on whole brain volume

We used a model-comparison approach to select the best combination of socio-ecological variables accounting for the variability in whole brain size across species using PGLS. Details of the different models are provided in supplementary material (SI section 3).

Earlier version: The best model explaining the volume of the whole brain is the one that includes body mass, daily traveled distance, and population density. The values of the estimated coefficients are provided in table 1. These three variables have a positive influence on the volume of the whole brain. Even if all three factors had a significant effect, their influence on the whole brain volume seemed to differ: there was an order of magnitude between the estimated weight (i.e. contribution) of body mass (0.47 +/- 0.15) and that of daily traveled distance (0.05 +/- 0.01). There was also an order of magnitude between the estimated weight of daily traveled distance and population density (0.007 +/-0.003).

Revised version: The best model explaining the volume of the whole brain is the one that includes body mass, daily traveled distance, and population density. The values of the estimated coefficients are provided in table 1. These three variables have a positive influence on the volume of the whole brain. Even if all three factors had a significant effect, their influence on the whole brain volume seemed to differ: there was an order of magnitude between the estimated beta (i.e. contribution) of body mass (0.47 +/-0.15) and that of daily traveled distance (0.05 +/-0.01). There was also an order of magnitude between the estimated beta of daily traveled distance and population density (0.007 +/-0.003). Note that the standardized beta (see methods) is also greater for daily traveled distance than for population density (see supplementary material section 5).

Earlier version: We evaluated the robustness of the influence of these socio-ecological variables on whole brain volume using a « leave-one-out » procedure. Thus, each model was tested 16 times (one per species removed) and we evaluated the reliability of each variable after each species was removed. The details of this analysis are provided in the supplementary material (SI section 4). The influence of body mass was very robust, in that it remained significant for all models where one species was removed. Daily traveled distance remained significant for all models but one, whereas population density lost significance (p>0.05) for four models where one species was removed, even if in all cases the estimated coefficient remained positive. Thus, the influence of body mass and daily traveled distance are very robust for the combination of species used for the analysis. It is less the case for population density, even if based on model comparison this variable plays a significant role in increasing the volume of the whole brain.

Revised version: We evaluated the robustness of the influence of these socio-ecological variables on whole brain volume using a « leave-one-out » procedure. Thus, each model was tested 16 times (one per species removed) and we evaluated the reliability of each variable after each species was removed. The details of this analysis are provided in the supplementary material (section 4). The influence of body mass was very robust, in that it remained significant for all models where one species was removed. Daily traveled distance remained significant for all models but one, whereas population density lost significance (p>0.05) for four models where one species was removed, even if in all cases the estimated coefficient remained positive. Thus, the influence of body mass and daily traveled distance are very robust for the combination of species used for the analysis. It is less the case for population density, even if based on model comparison this variable plays a significant role in increasing the volume of the whole brain.

Earlier version: In summary, this analysis indicates that the volume of the whole brain across primates is positively modulated by body mass, as well as variables related to both social (population density) and ecological (daily traveled distance) variables. The weight of the ecological variable, however, seems to be stronger and more reliable than the weight of the social variable.

Revised version: In summary, this analysis indicates that the volume of the whole brain across primates is positively modulated by body mass, as well as social (population density) and ecological (daily traveled distance) variables. The weight of the ecological variable, however, seems to be slightly stronger and more reliable than the weight of the social variable.

Influence of socio-ecological variables on Frontal Pole volume

Earlier version: As for the whole brain, the best model accounting for the volume of the frontal pole is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (SI section 3). The values of estimated coefficients, all positive, are provided in the table 2.

Revised version: As for the whole brain, the best model accounting for the volume of the frontal pole is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (section 3). The values of estimated coefficients, all positive, are provided in table 2.

Earlier version: As for the whole brain, the largest effect is body mass (0.69 +/- 0.23), then daily traveled distance (0.07 +/- 0.02) and population density (0.011 +/- 0.004). By comparison with the whole brain, however, the coefficients for daily traveled distance and population density have the same order of magnitude such that their impact on the volume of the frontal pole seems to be equivalent.

Revised version: As for the whole brain, the largest effect is body mass (0.69 +/- 0.23), then daily traveled distance (0.07 +/-0.02) and population density (0.011 +/-0.004). By comparison with the whole brain, however, the coefficients for daily traveled distance and population density have the same order of magnitude such that their impact on the volume of the frontal pole seems to be equivalent. Note that the standardized beta for daily traveled distance is smaller than for population density, whereas it was the opposite for the whole brain, i.e. bigger for daily traveled distance compared to population density.

Added in the revised version: As indicated in the methods section, we also conducted an analysis using a less conservative evaluation of FP volumes in monkeys. As shown in supplementary material (section 4), the relative weight of population density decreased when adding cortical regions dorsal to the principal sulcus. Thus, when using a less conservative estimate of the FP volume, its relation with socio-ecological variables was closer to that of the whole brain.

Earlier version: As for the whole brain, we evaluated the reliability of the model’s variables using a “leave-one-out” procedure at the species level. The results of this analysis were similar to what we reported for the whole brain: the influence of body mass remained significant whatever the species removed, only one species caused daily travelled distance to lose significance and three species caused population density to loose significance (p>0.05). As can be seen in the supplementary material (SI section 4), however, the decrease in the weight of the estimated coefficient for population density was less dramatic for these three species, compared to what we observed for the whole brain.

Revised version: As for the whole brain, we evaluated the reliability of the model’s variables using a ‘leave-one-out’ procedure at the species level. The results of this analysis were similar to what we reported for the whole brain: the influence of body mass remained significant whatever the species removed, only one species caused daily travelled distance to lose significance and three species caused population density to lose significance (p>0.05). As can be seen in the supplementary material (section 4), however, the decrease in the weight of the estimated coefficient for population density was less dramatic for these three species, compared to what we observed for the whole brain. As specified in the method section, we also conducted an analysis where the dorsal limit of the frontal pole was less conservative, to include a part of the brain which might belong to the FP but also to the DLPFC (see results in supplementary material, section 4).

Altogether, these data indicate that the volume of the frontal pole across primates is positively modulated by body mass as well as by both population density and daily traveled distance, in line with the idea that it is affected by both social and ecological components. By comparison with the whole brain, the relative influence of the social (Population density) and ecological (Daily Traveled Distance) variables are more balanced.

Influence of socio-ecological variables on Dorso-Lateral Prefrontal Cortex volume

Earlier version: As for the whole brain and the frontal pole, the best model for the DLPFC is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (SI section 3). The values of estimated coefficients, all positive, are provided in the table 3.

Revised version: As for the whole brain and the frontal pole, the best model for the DLPFC is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (section 3). The values of estimated coefficients, all positive, are provided in table 3.

Removed in the revised version: As for the whole brain and the frontal pole, the largest effect is body mass (0.47 +/- 0.18), followed by daily traveled distance (0.047 +/- 0.01). Even if the best model for DLPFC includes population density, its influence was very small (0.006 +/- 0.003) and failed to reach significance (p=0.067).

Earlier version: The « leave-one-out » procedure confirmed the reliability of the model for body mass (with no species causing that variable to lose significance when removed) and that of daily traveled distance (only one species causing that variable to loose significance when removed). The influence of population density only reached significance for two of the models where one species was removed, but it failed to reach significance for all other combinations tested (n=14), as it was the case for the original one with all 16 species. Details of this analysis can be found in the supplementary material, SI section 4. In other words, the relative weakness of the influence of population density compared to that of other variables is relatively reliable across the combinations of species used for the tests.

Revised version: As for the whole brain and the frontal pole, the largest effect is body mass (0.47 +/-0.18), followed by daily traveled distance (0.047 +/-0.01). Even if the best model for DLPFC includes population density, its influence was very small (0.006 +/-0.003) and failed to reach significance (p=0.067). The « leave-one-out » procedure confirmed the reliability of the model for body mass (with no species causing that variable to lose significance when removed) and that of daily traveled distance (only one species causing that variable to lose significance when removed). The influence of population density only reached significance for two of the models where one species was removed, but it failed to reach significance for all other combinations tested (n=14), as it was the case for the original one with all 16 species. Details of this analysis can be found in the supplementary material, section 4. In other words, the relative weakness of the influence of population density compared to that of other variables is relatively reliable across the combinations of species used for the tests.

In summary, the volume of the DLPFC across primates is also positively modulated by body mass as well as by both ecological (daily traveled distance) and social (population density) variables, but the influence of the later appears much weaker than that of the former.

Discussion

Earlier version: We evaluated the influence of eleven socio-ecological variables on the size of the whole brain as well as two specific brain regions involved in executive functions: the FP and the DLPFC. All together, our results showed that all three cerebral measures strongly correlated and were influenced by the same set of socio-ecological variables: body mass, daily traveled distance and population density. As expected from evidence suggesting brain-body covariation (Martin, 1981), body mass had a strong and reliable influence on all brain measures. Daily traveled distance, a proxy for how challenging is foraging, also had a clear positive influence on all brain measures. Finally, population density, a proxy for how challenging and complex are social interactions, was generally less powerful at explaining neuro-anatomical variability, but it was also the variable showing the greatest difference in effect size across brain regions. For the FP, the influence of population density was similar to that of daily travelled distance, but it was much weaker for the whole brain, and it even failed to reach significance for the DLPFC. Thus, our data are generally compatible with the idea that the evolution of executive functions relying upon FP and DLPFC is driven both by ecological and social constraints. Critically, however, the relative influence of these constraints seems to vary across regions, in line with our hypothesis based on their known cognitive functions in laboratory conditions.

Revised version: We evaluated the influence of eleven socio-ecological variables on the size of the whole brain as well as two specific brain regions involved in executive functions: the frontal pole and the dorsolateral prefrontal cortex. All together, our results showed that all three cerebral measures strongly correlated and were influenced by the same set of socio-ecological variables: body mass, daily traveled distance and population density. As expected from evidence suggesting brain-body covariation (Martin 1981), body mass had a strong and reliable influence on all brain measures. Daily traveled distance, a proxy for how challenging is foraging, also had a clear positive influence on all brain measures. Finally, population density, a proxy for how challenging and complex are social interactions, was generally less powerful at explaining neuro-anatomical variability, but it was also the variable showing the greatest difference in effect size across brain regions. For the FP, the influence of population density was similar to that of daily traveled distance, but it was much weaker for the whole brain, and it even failed to reach significance for the DLPFC. Thus, our data are generally compatible with the idea that the evolution of executive functions relying upon FP and DLPFC is driven both by ecological and social constraints. Critically, however, the relative influence of these constraints seems to vary across regions, in line with our hypothesis based on the cognitive functions with which they are associated in laboratory conditions. Thus, neuro-cognitive entities as defined in laboratory conditions can readily be articulated with socio-ecological processes in the wild, in order to better understand primate cognition and behavior in natural conditions.

First, we need to consider two methodological questions:

Earlier version: How reliable is the difference between brain regions? Even if the same combination of variables accounted for the relative size of all three brain measures, there was a difference in the weight of population density, proxy for social challenges, across the three regions. Indeed, for the FP, the effect of population density was as strong as that of daily traveled distance: both were significant, and their effect size (beta weight) was of the same order of magnitude. Also, they showed a similarly small sensitivity to the « leave-one-out » procedure. By contrast, for the DLPFC, the effect of population density failed to reach significance (p=0.07) and its beta weight was one order of magnitude smaller than that of daily traveled distance. Finally, this relatively weak influence was confirmed by the “leave-one-out” procedure, since the effect of population density on the size of the DLPFC was reliably marginal across all combination of species used to fit the model. Interestingly, the whole brain volume (which is strongly correlated with both FP and DLPFC) seems to show an intermediate tendency: even if the effect of population density appears smaller than that of daily traveled distance (beta weight is one order of magnitude smaller), it is clearly significant, and the « leave-one-out » procedure indicated that it was only slightly less reliable than daily traveled distance, with a significant decrease in model fit quality for four vs. one combination of species. Altogether, this indicates that even if the influence of population density cannot be ruled out for any of the three brain measures, it appears quantitatively smaller for the DLPFC than for the FP. This is compatible with the fact that the influence of population density seems to be of intermediate magnitude in the whole brain, which includes both FP and DLPFC.

Revised version: How reliable is the difference we reported between brain regions? Even if the same combination of variables accounted for the relative size of all three brain measures, there was a difference in the weight of population density, proxy for social challenges, across the three regions. Indeed, for the FP, the effect of population density was as strong as that of daily traveled distance: both were significant, and their effect size (beta weight) was of the same order of magnitude. Also, they showed a similarly small sensitivity to the « leave-one-out » procedure. Finally, standardized beta weights for population density were greater than for daily traveled distance. By contrast, for the DLPFC, the effect of population density failed to reach significance (p=0.07), its beta weight was one order of magnitude smaller than that of daily traveled distance and standardized beta weights for population density were smaller than for daily traveled distance. Finally, this relatively weak influence of population density on DLPFC volume was confirmed by the ‘leave-one-out’ procedure, since the effect of population density on the size of the DLPFC was reliably marginal across all combination of species used to fit the model. Interestingly, the whole brain volume (which is strongly correlated with both FP and DLPFC) seems to show an intermediate tendency: even if the effect of population density appears smaller than that of daily traveled distance (beta weight is one order of magnitude smaller), it is clearly significant, and the « leave-one-out » procedure indicated that it was only slightly less reliable than daily traveled distance, with a significant decrease in model fit quality for four vs. one combination of species. Altogether, this indicates that even if the influence of population density cannot be ruled out for any of the three brain measures, it appears quantitatively smaller for the DLPFC than for the FP. Note that the relative weight of population density also decreased (relative to that of daily traveled distance) when we used a less conservative estimate of FP volumes, which included a dorsal part potentially overlapping with the DLPFC. This is also compatible with the fact that the influence of population density seems to be of intermediate magnitude in the whole brain, which includes both FP and DLPFC.

Earlier version: To what extent do these conclusions depend upon the specific sample of species used here? The results of the « leave-one-out » analysis showed that our sample, even if relatively limited, was sufficient to evaluate the relative weight of socio-ecological variables on specific brain regions in primates. Interestingly, the relation between brain region volumes and socio-ecological variables was not completely independent from the sample of species, as previously shown in a study conducted on a larger sample (Powell et al., 2017). In line with the main analysis (PGLS with all species included), the influence of body mass and daily traveled distance appears very reliable (no more than one species caused the model to fail when removed). The influence of population density appeared less reliable, since it was sensitive to the removal of up to three-four species (as a function of brain region, see previous paragraph). But the influence of individual species is nearly impossible to interpret, because no apparent obvious pattern emerges. Indeed, the species without which the model failed was not always the same (e.g. humans and baboons have very distinct values for population density, 56 vs. 7.5). Note, however, that getting clear intuitions from the data remains difficult because the PGLS includes the covariance induced by phylogeny to the linear relation across variables. We see no reason to exclude humans (as done in previous studies e.g. (Alex R. DeCasien & Higham, 2019b), because even if the brain of Homo sapiens is much bigger than others, a general biological law describing the relation between socio-ecology and the brain in primates should apply to humans (Gabi et al., 2016; Herculano-Houzel, 2017).Typically, modern human populations show a tremendous geographical variability in terms of socio-ecological variables, but the same problematic applies to other modern primate species (including those confronted to intense anthropization of their habitat) (Beauchamp & Cabana, 1990). Thus, the actual potential limitation (which concerns all species) is the reliability with which socio-ecological variables were estimated for each species, given the amount of intra-specific variability. Note, there is also a significant amount of intra-specific variability at the level of the brain, both in humans and non-human primates, but this intra-specific variability was shown to be negligible compared to inter-specific variability (Alex R. DeCasien & Higham, 2019a; Louail et al., 2019; Maguire et al., 2000; Testard et al., 2022). Thus, we acknowledge that intra-specific variability could introduce noise in the inter-specific relation between brain measures and socio-ecological variables, but the fact that clear relations could be established with our sample indicates that this source of intraspecific variability was limited enough relative to the inter-species relation between neuroanatomical measures and socio-ecological variables. In other words, the error with which these variables were estimated at the level of individual species was small enough to allow us to study the relation of interest across species, and from that perspective, there is no reason to exclude any species from that analysis and our sample appears reliable enough to characterize the relation between neuroanatomical features and socio-ecological variables across primates.

Revised version: To what extent do these conclusions depend upon the specific sample of species used here? The results of the « leave-one-out » analysis showed that our sample, even if relatively limited, was sufficient to evaluate the relative weight of socio-ecological variables on specific brain regions in primates. Interestingly, the relation between brain region volumes and socio-ecological variables was not completely independent from the sample of species, as previously shown in a study conducted on a larger sample (Powell, Isler et al. 2017). In line with the main analysis (PGLS with all species included), the influence of body mass and daily traveled distance appears very reliable (no more than one species caused the model to fail when removed). The influence of population density appeared less reliable, since it was sensitive to the removal of up to three-four species (as a function of brain region, see previous paragraph). But the influence of individual species is nearly impossible to interpret, because no apparent obvious pattern emerges. Indeed, the species without which the model failed was not always the same (e.g. humans and baboons have very distinct values for population density, 56 vs. 7.5). Note, however, that getting clear intuitions from the data remains difficult because the PGLS includes the covariance induced by phylogeny to the linear relation across variables. We see no reason to exclude humans (as done in previous studies e.g. (DeCasien and Higham 2019), because even if the brain of Homo sapiens is much bigger than others, a general biological law describing the relation between socio-ecology and the brain in primates should apply to humans (Gabi, Neves et al. 2016, Herculano-Houzel 2017). Moreover, even if the human brain is much bigger than that of other primates, the global organization of the frontal cortex appears qualitatively similar in humans and other primates (Sallet, Mars et al. 2013, Gabi, Neves et al. 2016, Herculano-Houzel, Kaas et al. 2016, Barrett, Dawson et al. 2020, Roumazeilles, Eichert et al. 2020). One potential issue with humans is the difficulty to evaluate socio-ecological variables, given that modern human populations show a tremendous geographical variability in terms of socio-ecological variables. But the same problematic applies to other modern primate species (including those confronted to intense anthropization of their habitat) (Beauchamp and Cabana 1990). Thus, the actual potential limitation (which concerns all species) is the reliability with which socio-ecological variables were estimated for each species, given the amount of intraspecific variability. Note, there is also a significant amount of intra-specific variability at the level of the brain, both in humans and non-human primates, but this intra-specific variability was shown to be negligible compared to inter-specific variability (Maguire, Gadian et al. 2000, DeCasien and Higham 2019, Louail, Gilissen et al. 2019, Testard, Brent et al. 2022). Thus, we acknowledge that intra-specific variability could introduce noise in the inter-specific relation between brain measures and socio-ecological variables, but the fact that clear relations could be established with our sample indicates that this source of intra-specific variability was limited enough relative to the interspecies relation between neuroanatomical measures and socio-ecological variables. In other words, the error with which these variables were estimated at the level of individual species was small enough to allow us to study the relation of interest across species, and from that perspective, there is no reason to exclude any species from that analysis and our sample appears reliable enough to characterize the relation between neuroanatomical features and socio-ecological variables across primates.

Earlier version: Our results provide an original insight on the existing debate on the source of inter-species variability in cognitive abilities, as measured using brain volumes, and on the main socio-ecological variables that affect the evolution of specific brain regions. Given that all three brain measures were highly correlated, it is not surprising that their variability across species is predicted by the same model, i.e. the same combination of socio-ecological variables. This is in line with the intuition that across primates the major difference in brain volumes is often accounted for by a difference in size of the prefrontal cortex (Passingham et al., 2012); but see (A. R. DeCasien et al., 2022). Importantly, more precise quantitative studies showed the fraction of neurons in the prefrontal cortex, relative to the whole brain, (8%) was well conserved across primates, such that the key difference across species is not the proportion of prefrontal neurons, but its absolute number (Herculano-Houzel, Collins, Wong, & Kaas, 2007). In other words, the strong development of executive functions in species with larger prefrontal cortices is related to an absolute increase in number of neurons, rather than in an increase in the ration between the number of neurons in the PFC vs the rest of the brain. In that frame, it seems more appropriate to use absolute measures rather than relative measures, to evaluate the weight of each brain region and its corresponding cognitive function on behavior. In addition, the exact boundary between FP and DLPFC as functional entities remains difficult to assess, even if the functions attributed to these two regions are clearly distinct (Boorman, Behrens, Woolrich, & Rushworth, 2009; Gallagher & Frith, 2003; Genovesio, Wise, & Passingham, 2013; Koechlin, 2016; Passingham & Sakai, 2004; Preuss & Wise, 2021). Our aim here was clearly not to provide a clear identification of anatomical boundaries across brain regions in individual species, as others have done using much finer neuro-anatomical methods. Such a fine neuro-anatomical characterization appears impossible to carry on for a sample size of species compatible with PGLS. Rather, as specified in the methods, we tried to maximize the reliability of the landmarks that could be identified in all species to compare the relative size of specific functional areas. As for our previous study using a similar approach, this method was reliable enough to capture meaningful and specific effects of interests (Louail et al., 2019). Thus, even if we make no strong claim regarding the specific boundaries of FP and DLPFC, the anatomo-functional difference between these regions, as characterized in a few species used in laboratory experiments, could reliably be extrapolated to other species (Amiez et al., 2019; Bludau et al., 2014; J. Sallet et al., 2013). As we and others have done before (within and across species), we assume that the size of a given brain region provides a good proxy of the strength of its influence (through its known function) on behavior (Barton & Harvey, 2000; Alex R. DeCasien & Higham, 2019b; Louail et al., 2019; Maguire et al., 2000). Again, we assume that the size of a brain region is a good proxy for the number of neurons allocated to a given function (given the macroscopic connectivity pattern of the region and its local anatomo-functional features), given the known positive relaiton between number of cortical neurons and cognitive skills in primates (Herculano-Houzel, 2017, 2018). In that frame, the difference in size of these brain regions across species could readily be used as proxies for the relative importance of their corresponding functions (namely, metacognition for the FP and working memory for the DLPFC). Again, we clearly make no strong claim regarding the exact boundaries of the regions of interest, and their relations with existing laboratory studies in humans or macaques. Rather, we wanted to assess the relation between neuro-cognitive entities (e.g. FP/ metacognition vs. DLPFC/working memory) and socio-ecological variables, based on the assumption that the more a given cognitive function would be required to face a given socio-ecological challenge in a given species, the bigger the corresponding brain region would be. From that perspective, these data provide a new insight into the function of the FP and the DLPFC in primates.

Revised version: Our results provide an original insight on the existing debate on the source of inter-species variability in cognitive abilities, as measured using brain volumes, and on the main socio-ecological variables that affect the evolution of specific brain regions. Given that all three brain measures were highly correlated, it is not surprising that their variability across species is predicted by the same model, i.e. the same combination of socio-ecological variables. This is in line with the intuition that across primates the major difference in brain volumes is often accounted for by a difference in size of the prefrontal cortex (Passingham, Passingham et al. 2012); but see (DeCasien, Barton et al. 2022). Importantly, more precise quantitative studies showed that the fraction of neurons in the prefrontal cortex (PFC), relative to the whole brain (8%) was well conserved across primates, such that the key difference across species is not the proportion of prefrontal neurons, but its absolute number (Herculano-Houzel, Collins et al. 2007). In other words, the strong development of executive functions in species with larger prefrontal cortices is related to an absolute increase in number of neurons in the PFC, rather than in an increase in the ratio between the number of neurons in the PFC vs the rest of the brain. In that frame, it seems more appropriate to use absolute measures rather than relative measures, to evaluate the weight of each brain region and its corresponding cognitive function on behavior.

Added in the revised version: In addition, the exact boundary between FP and DLPFC as functional entities remains difficult to assess, even if the functions attributed to these two regions are clearly distinct (Gallagher and Frith 2003, Passingham and Sakai 2004, Boorman, Behrens et al. 2009, Genovesio, Wise et al. 2013, Koechlin 2016, Preuss and Wise 2021). Our aim here was clearly not to provide a clear identification of anatomical boundaries across brain regions in individual species, as others have done using much finer neuro-anatomical methods. Such a fine neuro-anatomical characterization appears impossible to carry on for a sample size of species compatible with PGLS. More importantly, we make no claim about specific mapping between the regions measured here and either the specific cognitive operations of interest (metacognition vs working memory) or the socio-ecological variables (e.g. social or foraging complexity). First, we are aware of the complexity of identifying neural processes underlying such high level cognitive functions, and even in well controlled laboratory conditions the exact boundaries of cortical regions remains challenging (Petrides 2005, Sallet, Mars et al. 2013, Mansouri, Koechlin et al. 2017). Rather, as specified in the methods, we tried to maximize the reliability of the landmarks that could be identified in all species to compare the relative size of areas that are essentially used as proxies for their known function, as defined in laboratory conditions in macaques and humans. The underlying assumption that we used was that the functional anatomy of the prefrontal cortex was conserved enough in primates such that traits that were common between humans and macaques were also shared with other primates (Sallet, Mars et al. 2013, Bludau, Eickhoff et al. 2014, Amiez, Sallet et al. 2019). Note that studies conducted in marmosets confirm that the functional anatomy of the primate prefrontal cortex is relatively well conserved (Dias, Robbins et al. 1996, Roberts, Tomic et al. 2007). Of course, the neural mechanisms underlying metacognition and working memory are extremely complicated and they involve myriads of other brain regions (e.g. temporal and parietal cortices, as well as subcortical systems), but the contribution of the FP to metacognition and of the DLPFC to working memory is so strong and reliable in humans and laboratory primates that it seems reasonable to assume that, given the relative level of conservation of prefrontal cortex anatomical organization across primates, FP and DLPFC should also play a key role in metacognition and working memory in other primate species, at least those studied here. As for our previous study using a similar approach, this method was reliable enough to capture meaningful and specific effects of interests (Louail, Gilissen et al. 2019). Thus, even if we make no strong claim regarding the specific boundaries of FP and DLPFC, the anatomo-functional difference between these regions, as characterized in a few species used in laboratory experiments is reliable enough to be captured by the variability in socio-ecological variables in a larger set of primates.

Added in the revised version: In summary, we are aware of the limitations of our study in terms of anatomical specificity, but these measures appear robust enough to draw original conclusions regarding the relation between functions attributed to FP and DLPFC (metacognition and working memory) in laboratory conditions and socio-ecological variables across 16 primate species. As we and others have done before (within and across species), we assume that the size of a given brain region provides a good proxy of the strength of its influence (through its known function) on behavior (Barton and Harvey 2000, Maguire, Gadian et al. 2000, DeCasien and Higham 2019, Louail, Gilissen et al. 2019). Again, we assume that the size of a brain region is a good proxy for the number of neurons allocated to its function, given the known positive relation between number of cortical neurons and cognitive skills in primates (Herculano-Houzel 2017, Herculano-Houzel 2018). Here, we restricted our analysis to a set of specific functional regions, namely the FP and the DLPFC, which implies that a given region and the neurons that it contains are characterized by a relatively homogenous connectivity pattern and physiological features, as demonstrated by converging evidences from anatomy, physiology, imaging and behavioral tests in laboratory animals (see methods for further details). In that frame, the difference in size of these brain regions across species could readily be used as proxies for the relative importance of their corresponding functions (namely, metacognition for the FP and working memory for the DLPFC). Again, we clearly make no strong claim regarding the exact boundaries of the regions of interest, and their relations with existing laboratory studies in humans or macaques. Rather, we wanted to assess the relation between neuro-cognitive entities (e.g. FP/metacognition vs. DLPFC/working memory) and socio-ecological variables, based on the assumption that the more a given cognitive function would be required to face a given socio-ecolo-gical challenge in a given species, the bigger the corresponding brain region would be. From that perspective, these data provide a new insight into the function of the FP and the DLPFC in primates.

Earlier version: The stronger sensitivity of the frontal pole to the variable ‘population density’ is reminiscent of the social brain hypothesis, i.e. the idea that social challenges favored the evolution of larger brains, and especially larger neocortex size, with species living in larger groups having bigger brains to deal with the associated complex social interactions (Dunbar, 1998). Critically, the size of the frontal pole and the corresponding development of metacognitive skills is not exclusively related to population density (proxy for social interactions) but also to daily traveled distance (proxy for foraging complexity). Thus, these data are also compatible with the ecological brain hypothesis (Milton, 1981). This dual relation between frontal pole size, social and ecological constraints might be accounted for by the fact that in the wild, social and ecological constraints remain strongly related, even if we did not find strong correlations between those factors in our data set. In reality, it is difficult to treat ecological and social factors as if they were disconnected (Henke-von der Malsburg, Kappeler, & Fichtel, 2020), and cognitive abilities associated with foraging are likely to play a role in social foraging tactics too (Street, Navarrete, Reader, & Laland, 2017). Indeed, greater population density implies a greater inter-individual competition for food, and thus potentially increases in daily traveled distances, and/or the development of sophisticated foraging skills to deal with the increase in scramble competition. But from a cognitive point of view, this also implies that the benefits associated with increased FP volume, and presumably an increase in metacognitive skills, could be related to both social and ecological functions. Indeed, as pointed out earlier, the development of executive functions in general, and metacognitive skills in particular, could have been a critical leverage to allow the development of both complex social interactions (through the use of Theory of Mind) and complex foraging (through flexible, context-dependent planning) (Garcia et al., 2021). Thus, our work indicates not only that the notions of ‘social brain’ and ‘ecological brain’ should not be mutually exclusive, even when considering specific brain regions. Rather, we confirm intuitions based on laboratory studies that the frontal pole, through its role in metacognition, belongs to both sets of brain regions defined as the social brain and the ecological brain, respectively. Further work would be needed to capture the metacognitive processes underlying social and foraging processes in the wild, as well as their interactions.

Revised version: The stronger sensitivity of the frontal pole to the variable ‘population density’ is reminiscent of the social brain hypothesis, i.e. the idea that social challenges favored the evolution of larger brains, and especially larger neocortex size, with species living in larger groups having bigger brains to deal with the associated complex social interactions (Dunbar 1998). Critically, the size of the frontal pole and the corresponding development of metacognitive skills is not exclusively related to population density (proxy for social interactions) but also to daily traveled distance (proxy for foraging complexity). Thus, these data are also compatible with the ecological brain hypothesis (Milton 1981). This dual relation between frontal pole size, social and ecological constraints might be accounted for by the fact that in the wild, social and ecological constraints remain strongly related, even if we did not find strong correlations between those factors in our data set. In reality, it is difficult to treat ecological and social factors as if they were disconnected (Henke-von der Malsburg, Kappeler et al. 2020), and cognitive abilities associated with foraging are likely to play a role in social foraging tactics too (Street, Navarrete et al. 2017). Indeed, greater population density implies a greater inter-individual competition for food, and thus potentially increases in daily traveled distances, and/or the development of sophisticated foraging skills to deal with the increase in scramble competition. But from a cognitive point of view, this also implies that the benefits associated with increased FP volume, and presumably an increase in metacognitive skills, could be related to both social and ecological functions. Indeed, as pointed out earlier, the development of executive functions in general, and metacognitive skills in particular, could have been a critical leverage to allow the development of both complex social interactions (through the use of Theory of Mind) and complex foraging (through flexible, context-dependent planning) (Garcia, Bouret et al. 2021, Shultz and Dunbar 2022). Thus, our work indicates that the notions of ‘social brain’ and ‘ecological brain’ should not be mutually exclusive, even when considering specific brain regions. Rather, we confirm intuitions based on laboratory studies that the frontal pole, through its role in metacognition, belongs to both sets of brain regions defined as the social brain and the ecological brain, respectively. Further work would be needed to capture the metacognitive processes underlying social and foraging processes in the wild, as well as their interactions.

Earlier version: The size of the DLPFC, which we used as a proxy for working memory and planning, showed a significant relation with daily traveled distance (proxy for foraging complexity), and to a lesser extent with population density (proxy for social complexity). This suggests that the cognitive functions at play in the DLPFC, i.e. working memory and planning, are critical for foraging in primates, and even more so when they need to travel long distances, and presumably have to deal with more complex navigation strategies. This is clearly in line with the global idea of the ecological brain hypothesis, but here we provide a critical insight into the specific neuro-cognitive operations associated with foraging strategies. The weaker influence of population density, a marker of social complexity, might be surprising at first sight, because a priori working memory and planning could also be strongly involved in complex social interactions (Garcia et al., 2021). But it is in line with laboratory data showing that social interactions seem to rely much more upon rostro-medial prefrontal regions compared to DLPFC (Fleming & Dolan, 2012; Frith, 2007; Jérôme Sallet et al., 2011; Testard et al., 2022). In other words, this work confirms the specificity of the ‘social brain’’ as defined in laboratory conditions as a set of structures specifically involved in social interactions (Frith, 2007; Rushworth et al., 2013). Critically, again, social and ecological functions are tightly intermingled in primates’ natural environment and more specific studies would be needed to clarify how DLPFC-related functions such as working memory are involved in natural conditions, when animals need to face both ecological and social challenges.

Revised version: The size of the DLPFC, which we used as a proxy for working memory and planning skills, showed a significant relation with daily traveled distance (proxy for foraging complexity), and to a lesser extent with population density (proxy for social complexity). This suggests that the cognitive functions at play in the DLPFC, i.e. working memory and planning, are critical for foraging in primates, especially when they need to travel long distances, and presumably have to deal with more complex navigation strategies. This is clearly in line with the global idea of the ecological brain hypothesis, but here we provide a critical insight into the specific neuro-cognitive operations associated with foraging strategies. The weaker influence of population density, a marker of social complexity, might be surprising at first sight, because a priori working memory and planning could also be strongly involved in complex social interactions (Garcia, Bouret et al. 2021). But it is in line with laboratory data showing that social interactions seem to rely much more upon rostro-medial prefrontal regions compared to DLPFC (Frith 2007, Sallet, Mars et al. 2011, Fleming and Dolan 2012, Testard, Brent et al. 2022). In other words, this work confirms the specificity of the ‘social brain’’ as defined in laboratory conditions as a set of structures specifically involved in social interactions (Frith 2007, Rushworth, Mars et al. 2013). Critically, again, social and ecological functions are tightly intermingled in primates’ natural environment and more specific studies would be needed to clarify how DLPFC-related functions such as working memory are involved in natural conditions, when animals need to face both ecological and social challenges.

Earlier version: This new set of results can be integrated in the general framework of the primate mosaic brain evolution, i.e. the different distinct structures varying in size both within and between species, and reflecting selection for cognitive skills. In line with recently published papers (A. R. DeCasien et al., 2022; Smaers et al., 2021; Carel P. van Schaik et al., 2021), we argue that the relation between ecology, neurobiology and cognition can be better captured with more specific brain measures, and more specific cognitive operations, than by using the traditional measures of relative brain size alone. We believe that comparisons between brain regions have the potential to identify which patterns of brain region evolution can explain the evolution of brain size and then provide insights into the evolution of cognition. Our new set of results are in agreement with a previous study we conducted on another brain region, the ventromedial prefrontal cortex (VMPFC), critically involved in value-based decision making (Louail et al., 2019). This study was conducted on 29 brain scans from only five species, such that we could not use a PGLS and instead included phylogenetic distance as a co-regressor, along social and ecological variables. As we did here, we also identified the combination of socio-ecological variables that best predicted neuro-anatomical variability across species. Interestingly, the pattern reported for the whole brain was similar to that of the current study and, the weight of the ecological variable (daily traveled distance) was one order of magnitude greater than the influence of the social variable (group size). As it was the case here, VMPFC and whole brain volumes were strongly correlated (r=0.99), but the size of the VMPFC was predicted by a distinct set of ecological variables (dietary quality and weaning age), with little modulation by group size. Thus, our results from these two combined studies suggest specific neuro-cognitive entities, established in laboratory studies and defined by a conjunction of specific brain regions and cognitive operations, can be related to specific socio-ecological challenges that animals face in natural conditions.

Revised version: This new set of results can be integrated in the general framework of the primate mosaic brain evolution, i.e. the different distinct structures varying in size both within and between species, and reflecting selection for cognitive skills. In line with recently published papers (Smaers, Rothman et al. 2021, van Schaik, Triki et al. 2021, DeCasien, Barton et al. 2022), we argue that the relation between ecology, neurobiology and cognition can be better captured with more specific brain measures, and more specific cognitive operations, than by using the traditional measures of relative brain size alone. We believe that comparisons between brain regions have the potential to identify which patterns of brain region evolution can provide insights into the evolution of cognition. Our new set of results are in agreement with a previous study we conducted on another brain region, the ventromedial prefrontal cortex (VMPFC), critically involved in value-based decision making (Louail, Gilissen et al. 2019). This study was conducted on 29 brain scans from only five species, such that we could not use a PGLS and instead included phylogenetic distance as a co-regressor, along social and ecological variables. As we did here, we also identified the combination of socioecological variables that best predicted neuro-anatomical variability across species. Interestingly, the pattern reported for the whole brain was similar to that of the current study and, the weight of the ecological variable (daily traveled distance) was one order of magnitude greater than the influence of the social variable (group size). As it was the case here, VMPFC and whole brain volumes were strongly correlated (r=0.99), but the size of the VMPFC was predicted by a distinct set of ecological variables (dietary quality and weaning age), with little modulation by group size. Thus, our results from these two combined studies suggest that specific neuro-cognitive entities, established in laboratory studies and defined by a conjunction of specific brain regions and cognitive operations, can be related to specific socio-ecological challenges that animals face in natural conditions.

Earlier version: In conclusion, our results confirm that the size of specific brain regions can be related to socio-ecological variables through the cognitive operations relying on these regions in laboratory conditions. Thus, our approach aiming at articulating cognitive operations reported in laboratory settings with real socio-ecological challenges should provide a clear insight into the neuro-cognitive operations at play in the wild, as well as their evolution in primates. Conversely, integrating realistic socio-ecological challenges can provide a strong insight into the evolution of specific brain functions in primates. Of course, we would need to provide a clearer model regarding how facing socio-ecological challenges can rely upon specific and dynamic cognitive operations, but we believe that this is a critical first step that helps building a common theoretical framework to cross boundaries across behavioral ecology and cognitive neurosciences.

Revised version: In conclusion, our results confirm that the size of specific brain regions can be related to socio-eco-logical variables through the cognitive operations relying on these regions in laboratory conditions. Thus, our approach aiming at articulating cognitive operations reported in laboratory settings with real socio-ecological challenges should provide a clear insight into the neuro-cognitive operations at play in the wild, as well as their evolution in primates. Conversely, integrating realistic socio-eco-logical challenges can provide a strong insight into the evolution of specific brain functions in primates. Of course, we would need to provide a clearer model regarding how facing socio-ecological challenges can rely upon specific and dynamic cognitive operations, but we believe that this is a critical first step that helps building a common theoretical framework to cross boundaries across behavioral ecology and cognitive neurosciences.

Added in the revised version: Supporting information